[1484370084] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.47%) too high
[1484370144] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 34.69Mb (1.73%), Average OUT: 54.66Mb (2.73%)Total RX: 177034.43Gbits, Total TX: 128579.61Gbits
[1484370624] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.56%) too high
[1484370685] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (120962256195.90%) too high
[1484370744] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 63.26Mb (3.16%), Average OUT: 74.67Mb (3.73%)Total RX: 177061.83Gbits, Total TX: 128618.36Gbits
[1484371105] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.41%) too high
[1484371164] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 86.04Mb (4.30%), Average OUT: 57.89Mb (2.89%)Total RX: 177088.08Gbits, Total TX: 128646.31Gbits
[1484371644] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.29%) too high
[1484371704] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 624.03Mb (31.20%), Average OUT: 79.87Mb (3.99%)Total RX: 177158.51Gbits, Total TX: 128684.58Gbits
[1484371764] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293798.67%) too high
[1484371824] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 47.47Mb (2.37%), Average OUT: 66.66Mb (3.33%)Total RX: 177130.09Gbits, Total TX: 128693.24Gbits
[1484372124] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.53%) too high
[1484372185] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 48.26Mb (2.41%), Average OUT: 70.89Mb (3.54%)Total RX: 177148.05Gbits, Total TX: 128719.43Gbits
[1484372605] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.03%) too high
[1484372664] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 63.23Mb (3.16%), Average OUT: 659.79Mb (32.99%)Total RX: 177175.86Gbits, Total TX: 128787.73Gbits
[1484373085] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.29%) too high
[1484373145] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 61.59Mb (3.08%), Average OUT: 61.86Mb (3.09%)Total RX: 177206.19Gbits, Total TX: 128787.28Gbits
[1484373504] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671661.73%) too high
[1484373565] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 87.27Mb (4.36%), Average OUT: 78.05Mb (3.90%)Total RX: 177237.59Gbits, Total TX: 128821.72Gbits
[1484373874] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484373937] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 508.88b (0.00%)Total RX: 17.62Kbits, Total TX: 489.58Mbits
[1484373985] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.08%) too high
[1484374045] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 71.57Mb (3.58%), Average OUT: 50.23Mb (2.51%)Total RX: 177275.12Gbits, Total TX: 128853.71Gbits
[1484374190] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1484374256] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1484374318] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1484374376] HOST ALERT: INASOL-STEVEN;DOWN;HARD;4;PING CRITICAL - Packet loss = 100%
[1484374604] HOST ALERT: INASOL-STEVEN;UP;HARD;4;PING OK - Packet loss = 70%, RTA = 7.09 ms
[1484375125] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.92%) too high
[1484375185] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 80.05Mb (4.00%), Average OUT: 58.14Mb (2.91%)Total RX: 177370.17Gbits, Total TX: 128922.89Gbits
[1484375665] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.68%) too high
[1484375723] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 108.10Mb (5.41%), Average OUT: 65.04Mb (3.25%)Total RX: 177412.99Gbits, Total TX: 128958.32Gbits
[1484375965] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (119011252065.26%) too high
[1484376024] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 101.71Mb (5.09%), Average OUT: 78.77Mb (3.94%)Total RX: 177445.05Gbits, Total TX: 129015.22Gbits
[1484376144] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.95%) too high
[1484376205] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 111.13Mb (5.56%), Average OUT: 61.70Mb (3.08%)Total RX: 177464.76Gbits, Total TX: 128992.32Gbits
[1484376624] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.13%) too high
[1484376684] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 154.15Mb (7.71%), Average OUT: 639.29Mb (31.96%)Total RX: 177521.43Gbits, Total TX: 129059.40Gbits
[1484376748] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1484376836] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1484376897] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1484376957] HOST ALERT: INASOL-STEVEN;DOWN;HARD;4;PING CRITICAL - Packet loss = 100%
[1484376972] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.40%) too high
[1484377033] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 67.37Mb (67.37%), Average OUT: 3.75Mb (3.75%)Total RX: 26436.41Gbits, Total TX: 6266.95Gbits
[1484377225] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.02%) too high
[1484377284] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 161.70Mb (8.08%), Average OUT: 58.69Mb (2.93%)Total RX: 177608.37Gbits, Total TX: 129059.78Gbits
[1484377573] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.72%) too high
[1484377634] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 56.63Mb (56.63%), Average OUT: 6.27Mb (6.27%)Total RX: 26467.88Gbits, Total TX: 6269.95Gbits
[1484377825] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.21%) too high
[1484377873] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.98%) too high
[1484377885] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 188.35Mb (9.42%), Average OUT: 83.41Mb (4.17%)Total RX: 177733.86Gbits, Total TX: 129098.41Gbits
[1484377933] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 73.88Mb (73.88%), Average OUT: 5.06Mb (5.06%)Total RX: 26486.91Gbits, Total TX: 6272.01Gbits
[1484378113] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (76.60%) too high
[1484378174] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 44.56Mb (44.56%), Average OUT: 5.21Mb (5.21%)Total RX: 26502.28Gbits, Total TX: 6273.21Gbits
[1484378305] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.61%) too high
[1484378365] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 300.62Mb (15.03%), Average OUT: 69.68Mb (3.48%)Total RX: 177848.35Gbits, Total TX: 129130.28Gbits
[1484378724] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671661.28%) too high
[1484378772] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.45%) too high
[1484378785] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 316.04Mb (15.80%), Average OUT: 84.16Mb (4.21%)Total RX: 177972.90Gbits, Total TX: 129164.56Gbits
[1484378833] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 63.51Mb (63.51%), Average OUT: 6.62Mb (6.62%)Total RX: 26539.13Gbits, Total TX: 6276.64Gbits
[1484379133] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (81.51%) too high
[1484379145] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.16%) too high
[1484379193] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (94.26%) too high
[1484379205] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 331.35Mb (16.57%), Average OUT: 68.68Mb (3.43%)Total RX: 178097.71Gbits, Total TX: 129197.54Gbits
[1484379253] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (94.25%) too high
[1484379312] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (95.75%) too high
[1484379511] HOST ALERT: INASOL-STEVEN;UP;HARD;4;PING OK - Packet loss = 0%, RTA = 6.50 ms
[1484379673] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (81.95%) too high
[1484379673] SERVICE ALERT: INASOL-BGP01;INASOL [INTERNET REDCA - G0/0/0];WARNING;SOFT;1;WARNING - OUT bandwidth (82.46%) too high
[1484379685] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.35%) too high
[1484379733] SERVICE ALERT: INASOL-BGP01;INASOL [INTERNET REDCA - G0/0/0];WARNING;SOFT;2;WARNING - OUT bandwidth (78.32%) too high
[1484379745] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 411.53Mb (20.58%), Average OUT: 63.28Mb (3.16%)Total RX: 178300.50Gbits, Total TX: 129231.93Gbits
[1484379793] SERVICE ALERT: INASOL-BGP01;INASOL [INTERNET REDCA - G0/0/0];OK;SOFT;3;OK - Average IN: 11.22Mb (11.22%), Average OUT: 14.48Mb (14.48%)Total RX: 23648.49Gbits, Total TX: 31880.40Gbits
[1484380093] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (97.14%) too high
[1484380285] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.20%) too high
[1484380345] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 409.77Mb (20.49%), Average OUT: 70.91Mb (3.55%)Total RX: 178536.43Gbits, Total TX: 129269.31Gbits
[1484380393] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (88.21%) too high
[1484380393] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1484380453] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (92.61%) too high
[1484380753] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (86.86%) too high
[1484380765] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.15%) too high
[1484380825] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 495.66Mb (24.78%), Average OUT: 65.33Mb (3.27%)Total RX: 178748.67Gbits, Total TX: 129300.01Gbits
[1484380993] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (92.66%) too high
[1484381413] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (83.87%) too high
[1484381532] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.78%) too high
[1484381713] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (88.41%) too high
[1484381773] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.06%) too high
[1484381784] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.37%) too high
[1484381833] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (87.43%) too high
[1484381845] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 438.10Mb (21.91%), Average OUT: 89.96Mb (4.50%)Total RX: 179210.43Gbits, Total TX: 129372.45Gbits
[1484382013] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 70.75Mb (70.75%), Average OUT: 14.48Mb (14.48%)Total RX: 26819.14Gbits, Total TX: 6321.19Gbits
[1484382145] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256215.09%) too high
[1484382193] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (77.18%) too high
[1484382205] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (122978293799.87%) too high
[1484382253] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 67.27Mb (67.27%), Average OUT: 14.29Mb (14.29%)Total RX: 26835.97Gbits, Total TX: 6325.08Gbits
[1484382265] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 463.29Mb (23.16%), Average OUT: 77.68Mb (3.88%)Total RX: 179408.69Gbits, Total TX: 129407.18Gbits
[1484382432] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (76.46%) too high
[1484382493] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (81.35%) too high
[1484382554] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (83.54%) too high
[1484382565] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293801.10%) too high
[1484382611] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (76.58%) too high
[1484382625] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 471.05Mb (23.55%), Average OUT: 660.73Mb (33.04%)Total RX: 179581.36Gbits, Total TX: 129474.40Gbits
[1484382793] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 74.56Mb (74.56%), Average OUT: 13.05Mb (13.05%)Total RX: 26878.08Gbits, Total TX: 6333.57Gbits
[1484383045] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.31%) too high
[1484383103] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 444.61Mb (22.23%), Average OUT: 86.66Mb (4.33%)Total RX: 179811.64Gbits, Total TX: 129477.26Gbits
[1484383153] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (77.67%) too high
[1484383213] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (78.15%) too high
[1484383272] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (78.95%) too high
[1484383333] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (82.61%) too high
[1484383405] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293801.29%) too high
[1484383453] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 71.72Mb (71.72%), Average OUT: 13.84Mb (13.84%)Total RX: 26927.22Gbits, Total TX: 6341.88Gbits
[1484383465] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 390.93Mb (19.55%), Average OUT: 92.44Mb (4.62%)Total RX: 179970.73Gbits, Total TX: 129509.58Gbits
[1484383513] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (80.97%) too high
[1484383573] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (76.02%) too high
[1484383633] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 69.80Mb (69.80%), Average OUT: 13.99Mb (13.99%)Total RX: 26940.83Gbits, Total TX: 6344.60Gbits
[1484383765] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293801.18%) too high
[1484383813] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (80.14%) too high
[1484383825] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 436.43Mb (21.82%), Average OUT: 91.01Mb (4.55%)Total RX: 180126.41Gbits, Total TX: 129545.31Gbits
[1484383872] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 74.99Mb (74.99%), Average OUT: 14.69Mb (14.69%)Total RX: 26958.87Gbits, Total TX: 6347.78Gbits
[1484383933] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (83.93%) too high
[1484383993] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (84.64%) too high
[1484384005] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293818.34%) too high
[1484384052] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (78.44%) too high
[1484384065] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 426.68Mb (21.33%), Average OUT: 99.98Mb (5.00%)Total RX: 180234.07Gbits, Total TX: 129602.68Gbits
[1484384085] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1484384113] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;4;OK - Average IN: 72.15Mb (72.15%), Average OUT: 12.78Mb (12.78%)Total RX: 26978.10Gbits, Total TX: 6350.61Gbits
[1484384125] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.66%) too high
[1484384158] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1484384185] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 394.52Mb (19.73%), Average OUT: 97.78Mb (4.89%)Total RX: 180283.39Gbits, Total TX: 129579.66Gbits
[1484384218] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;3;(Host check timed out after 31.00 seconds)
[1484384277] HOST ALERT: INASOL-STEVEN;DOWN;HARD;4;PING CRITICAL - Packet loss = 100%
[1484384353] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.91%) too high
[1484384413] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (80.21%) too high
[1484384473] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 70.01Mb (70.01%), Average OUT: 12.16Mb (12.16%)Total RX: 27004.01Gbits, Total TX: 6354.39Gbits
[1484384545] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.31%) too high
[1484384552] HOST ALERT: INASOL-STEVEN;UP;HARD;4;PING OK - Packet loss = 0%, RTA = 8.40 ms
[1484384605] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (122978293817.62%) too high
[1484384653] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (82.00%) too high
[1484384665] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 395.62Mb (19.78%), Average OUT: 679.71Mb (33.99%)Total RX: 180472.21Gbits, Total TX: 129654.73Gbits
[1484384713] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (87.30%) too high
[1484384773] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (82.72%) too high
[1484384832] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (81.58%) too high
[1484384904] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671661.23%) too high
[1484384965] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 993.39Mb (49.67%), Average OUT: 72.54Mb (3.63%)Total RX: 180628.56Gbits, Total TX: 129647.60Gbits
[1484384974] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484385025] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293815.39%) too high
[1484385034] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 474.24b (0.00%)Total RX: 17.62Kbits, Total TX: 495.37Mbits
[1484385085] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 373.29Mb (18.66%), Average OUT: 656.78Mb (32.84%)Total RX: 180639.75Gbits, Total TX: 129692.01Gbits
[1484385191] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (95.14%) too high
[1484385313] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (88.93%) too high
[1484385324] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.95%) too high
[1484385385] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 353.77Mb (17.69%), Average OUT: 71.02Mb (3.55%)Total RX: 180764.98Gbits, Total TX: 129681.69Gbits
[1484385673] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (91.98%) too high
[1484385793] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (89.89%) too high
[1484385985] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293818.30%) too high
[1484386044] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 408.53Mb (20.43%), Average OUT: 79.20Mb (3.96%)Total RX: 181022.80Gbits, Total TX: 129769.39Gbits
[1484386165] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.58%) too high
[1484386225] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 423.00Mb (21.15%), Average OUT: 86.15Mb (4.31%)Total RX: 181102.19Gbits, Total TX: 129750.88Gbits
[1484386333] SERVICE ALERT: INASOL-CyM_ABOGADOS;INASOL - CyM_ABOGADOS;CRITICAL;SOFT;1;CRITICAL - IN bandwidth (2459565875337.25%) too high
[1484386333] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (95.53%) too high
[1484386393] SERVICE ALERT: INASOL-CyM_ABOGADOS;INASOL - CyM_ABOGADOS;OK;SOFT;2;OK - Average IN: 140.97Kb (0.14%), Average OUT: 45.28Kb (0.05%)Total RX: 8.71Mbits, Total TX: 2.83Mbits
[1484386513] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (81.31%) too high
[1484386584] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671661.38%) too high
[1484386645] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 451.76Mb (22.59%), Average OUT: 88.56Mb (4.43%)Total RX: 181279.43Gbits, Total TX: 129786.00Gbits
[1484386932] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (92.18%) too high
[1484386993] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (81.31%) too high
[1484387053] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (95.30%) too high
[1484387173] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (86.74%) too high
[1484387233] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.05%) too high
[1484387293] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (88.72%) too high
[1484387365] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.91%) too high
[1484387424] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 408.99Mb (20.45%), Average OUT: 90.59Mb (4.53%)Total RX: 181582.10Gbits, Total TX: 129852.82Gbits
[1484387653] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (92.97%) too high
[1484387713] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (88.92%) too high
[1484387785] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.68%) too high
[1484387845] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 398.80Mb (19.94%), Average OUT: 90.25Mb (4.51%)Total RX: 181758.00Gbits, Total TX: 129888.60Gbits
[1484387893] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (91.42%) too high
[1484388013] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (89.03%) too high
[1484388145] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.88%) too high
[1484388205] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 381.09Mb (19.05%), Average OUT: 659.07Mb (32.95%)Total RX: 181896.07Gbits, Total TX: 129955.28Gbits
[1484388253] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.03%) too high
[1484388313] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (84.52%) too high
[1484388373] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (96.28%) too high
[1484388433] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (82.63%) too high
[1484388564] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.47%) too high
[1484388613] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 74.08Mb (74.08%), Average OUT: 13.64Mb (13.64%)Total RX: 27358.22Gbits, Total TX: 6408.13Gbits
[1484388625] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 353.32Mb (17.67%), Average OUT: 72.65Mb (3.63%)Total RX: 182042.09Gbits, Total TX: 129955.66Gbits
[1484388673] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.37%) too high
[1484388733] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 72.29Mb (72.29%), Average OUT: 14.91Mb (14.91%)Total RX: 27367.08Gbits, Total TX: 6409.94Gbits
[1484388865] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293815.94%) too high
[1484388873] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484388913] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (83.40%) too high
[1484388925] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 392.65Mb (19.63%), Average OUT: 651.89Mb (32.59%)Total RX: 182153.62Gbits, Total TX: 130011.47Gbits
[1484388935] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 486.8b (0.00%)Total RX: 17.62Kbits, Total TX: 497.41Mbits
[1484388973] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (77.76%) too high
[1484389033] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (77.33%) too high
[1484389045] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.84%) too high
[1484389093] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;4;OK - Average IN: 73.06Mb (73.06%), Average OUT: 14.21Mb (14.21%)Total RX: 27394.51Gbits, Total TX: 6414.98Gbits
[1484389104] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 411.37Mb (20.57%), Average OUT: 65.52Mb (3.28%)Total RX: 182227.74Gbits, Total TX: 129990.06Gbits
[1484389344] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671679.67%) too high
[1484389405] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 462.88Mb (23.14%), Average OUT: 72.77Mb (3.64%)Total RX: 182366.99Gbits, Total TX: 130046.98Gbits
[1484389525] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.38%) too high
[1484389583] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 406.20Mb (20.31%), Average OUT: 82.33Mb (4.12%)Total RX: 182445.77Gbits, Total TX: 130027.00Gbits
[1484389633] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (79.37%) too high
[1484389692] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (92.95%) too high
[1484389753] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (91.97%) too high
[1484389812] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.06%) too high
[1484389873] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (84.96%) too high
[1484389945] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.54%) too high
[1484390004] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 509.79Mb (25.49%), Average OUT: 77.00Mb (3.85%)Total RX: 182634.65Gbits, Total TX: 130058.48Gbits
[1484390173] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 73.79Mb (73.79%), Average OUT: 12.43Mb (12.43%)Total RX: 27477.55Gbits, Total TX: 6428.68Gbits
[1484390185] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256216.15%) too high
[1484390245] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 419.67Mb (20.98%), Average OUT: 85.09Mb (4.25%)Total RX: 182743.85Gbits, Total TX: 130078.51Gbits
[1484390353] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (83.14%) too high
[1484390413] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (80.82%) too high
[1484390473] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 59.73Mb (59.73%), Average OUT: 9.67Mb (9.67%)Total RX: 27499.53Gbits, Total TX: 6431.64Gbits
[1484390845] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.31%) too high
[1484390905] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 1.05Gb (52.32%), Average OUT: 88.72Mb (4.44%)Total RX: 183068.14Gbits, Total TX: 130128.68Gbits
[1484390965] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293822.63%) too high
[1484391025] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 512.75Mb (25.64%), Average OUT: 76.82Mb (3.84%)Total RX: 183096.38Gbits, Total TX: 130138.45Gbits
[1484391193] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1484391265] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.37%) too high
[1484391325] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 494.06Mb (24.70%), Average OUT: 78.96Mb (3.95%)Total RX: 183242.33Gbits, Total TX: 130161.60Gbits
[1484391372] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (82.40%) too high
[1484391433] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (77.58%) too high
[1484391493] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (75.66%) too high
[1484391553] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;4;OK - Average IN: 65.04Mb (65.04%), Average OUT: 11.01Mb (11.01%)Total RX: 27572.27Gbits, Total TX: 6444.92Gbits
[1484391685] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293801.06%) too high
[1484391745] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 419.13Mb (20.96%), Average OUT: 87.98Mb (4.40%)Total RX: 183440.26Gbits, Total TX: 130200.01Gbits
[1484391913] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (82.23%) too high
[1484391973] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (92.55%) too high
[1484391985] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256198.13%) too high
[1484392033] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (92.07%) too high
[1484392045] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 411.35Mb (20.57%), Average OUT: 95.05Mb (4.75%)Total RX: 183560.19Gbits, Total TX: 130229.29Gbits
[1484392093] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (94.54%) too high
[1484392153] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (88.13%) too high
[1484392164] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671676.13%) too high
[1484392225] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 347.23Mb (17.36%), Average OUT: 93.09Mb (4.65%)Total RX: 183627.20Gbits, Total TX: 130286.52Gbits
[1484392273] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (91.78%) too high
[1484392333] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (83.61%) too high
[1484392333] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1484392345] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.15%) too high
[1484392393] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 74.32Mb (74.32%), Average OUT: 10.33Mb (10.33%)Total RX: 27639.03Gbits, Total TX: 6454.82Gbits
[1484392405] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 376.52Mb (18.83%), Average OUT: 654.11Mb (32.71%)Total RX: 183694.57Gbits, Total TX: 130301.28Gbits
[1484392453] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.94%) too high
[1484392512] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 73.31Mb (73.31%), Average OUT: 11.19Mb (11.19%)Total RX: 27648.09Gbits, Total TX: 6456.17Gbits
[1484392753] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (80.22%) too high
[1484392765] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.53%) too high
[1484392774] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484392813] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (80.52%) too high
[1484392823] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 385.30Mb (19.26%), Average OUT: 64.66Mb (3.23%)Total RX: 183860.27Gbits, Total TX: 130296.75Gbits
[1484392833] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 455.2b (0.00%)Total RX: 17.62Kbits, Total TX: 499.44Mbits
[1484392873] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 74.02Mb (74.02%), Average OUT: 11.99Mb (11.99%)Total RX: 27674.60Gbits, Total TX: 6460.53Gbits
[1484393052] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.05%) too high
[1484393113] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 70.56Mb (70.56%), Average OUT: 10.09Mb (10.09%)Total RX: 27691.49Gbits, Total TX: 6463.09Gbits
[1484393245] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256212.42%) too high
[1484393293] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (80.90%) too high
[1484393304] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (125062671660.27%) too high
[1484393353] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 73.90Mb (73.90%), Average OUT: 14.11Mb (14.11%)Total RX: 27708.74Gbits, Total TX: 6466.22Gbits
[1484393364] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 348.55Mb (17.43%), Average OUT: 72.38Mb (3.62%)Total RX: 184061.29Gbits, Total TX: 130333.90Gbits
[1484393413] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (77.70%) too high
[1484393473] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 69.13Mb (69.13%), Average OUT: 11.59Mb (11.59%)Total RX: 27717.55Gbits, Total TX: 6467.66Gbits
[1484393612] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484393672] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 479.36b (0.00%)Total RX: 17.62Kbits, Total TX: 499.88Mbits
[1484393713] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (84.66%) too high
[1484393773] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 70.01Mb (70.01%), Average OUT: 11.55Mb (11.55%)Total RX: 27739.03Gbits, Total TX: 6471.18Gbits
[1484393845] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.93%) too high
[1484393905] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 337.19Mb (16.86%), Average OUT: 57.36Mb (2.87%)Total RX: 184246.46Gbits, Total TX: 130366.14Gbits
[1484393953] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1484394325] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.76%) too high
[1484394385] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 318.61Mb (15.93%), Average OUT: 83.98Mb (4.20%)Total RX: 184405.31Gbits, Total TX: 130401.07Gbits
[1484394973] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.34%) too high
[1484395031] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 74.04Mb (74.04%), Average OUT: 9.02Mb (9.02%)Total RX: 27817.71Gbits, Total TX: 6487.49Gbits
[1484395225] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.98%) too high
[1484395285] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 316.09Mb (15.80%), Average OUT: 73.10Mb (3.66%)Total RX: 184696.70Gbits, Total TX: 130469.62Gbits
[1484395393] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (76.47%) too high
[1484395453] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 69.80Mb (69.80%), Average OUT: 11.17Mb (11.17%)Total RX: 27847.60Gbits, Total TX: 6492.06Gbits
[1484395705] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.75%) too high
[1484395764] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 298.24Mb (14.91%), Average OUT: 63.18Mb (3.16%)Total RX: 184857.16Gbits, Total TX: 130504.43Gbits
[1484396052] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (84.52%) too high
[1484396113] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 69.07Mb (69.07%), Average OUT: 8.50Mb (8.50%)Total RX: 27887.89Gbits, Total TX: 6498.81Gbits
[1484396174] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (81.14%) too high
[1484396233] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 74.43Mb (74.43%), Average OUT: 10.21Mb (10.21%)Total RX: 27897.23Gbits, Total TX: 6500.04Gbits
[1484396245] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.24%) too high
[1484396293] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (77.34%) too high
[1484396305] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 324.06Mb (16.20%), Average OUT: 63.34Mb (3.17%)Total RX: 185025.48Gbits, Total TX: 130538.49Gbits
[1484396353] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (79.19%) too high
[1484396413] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 74.89Mb (74.89%), Average OUT: 10.80Mb (10.80%)Total RX: 27911.11Gbits, Total TX: 6501.95Gbits
[1484396473] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (82.93%) too high
[1484396533] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (89.35%) too high
[1484396592] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (97.58%) too high
[1484396653] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (92.86%) too high
[1484396712] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (83.21%) too high
[1484396725] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.90%) too high
[1484396785] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 298.49Mb (14.92%), Average OUT: 70.71Mb (3.54%)Total RX: 185170.02Gbits, Total TX: 130575.43Gbits
[1484396893] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 73.33Mb (73.33%), Average OUT: 11.71Mb (11.71%)Total RX: 27952.12Gbits, Total TX: 6507.76Gbits
[1484396954] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (86.91%) too high
[1484397013] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 65.81Mb (65.81%), Average OUT: 10.57Mb (10.57%)Total RX: 27961.30Gbits, Total TX: 6509.02Gbits
[1484397445] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293806.99%) too high
[1484397492] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (91.28%) too high
[1484397505] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 226.34Mb (11.32%), Average OUT: 67.07Mb (3.35%)Total RX: 185362.20Gbits, Total TX: 130658.01Gbits
[1484397554] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (78.33%) too high
[1484397612] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (80.06%) too high
[1484397673] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;4;OK - Average IN: 73.24Mb (73.24%), Average OUT: 17.06Mb (17.06%)Total RX: 28008.19Gbits, Total TX: 6517.17Gbits
[1484397745] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293807.54%) too high
[1484397804] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 231.22Mb (11.56%), Average OUT: 76.12Mb (3.81%)Total RX: 185432.27Gbits, Total TX: 130643.73Gbits
[1484398165] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.37%) too high
[1484398225] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 222.58Mb (11.13%), Average OUT: 62.29Mb (3.11%)Total RX: 185536.89Gbits, Total TX: 130676.45Gbits
[1484398573] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.51%) too high
[1484398585] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293811.30%) too high
[1484398633] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (75.04%) too high
[1484398645] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (122978293799.50%) too high
[1484398693] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 66.71Mb (66.71%), Average OUT: 14.56Mb (14.56%)Total RX: 28075.81Gbits, Total TX: 6530.99Gbits
[1484398705] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 258.86Mb (12.94%), Average OUT: 75.40Mb (3.77%)Total RX: 185662.70Gbits, Total TX: 130711.73Gbits
[1484398825] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293810.66%) too high
[1484398885] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 268.37Mb (13.42%), Average OUT: 82.51Mb (4.13%)Total RX: 185713.23Gbits, Total TX: 130758.62Gbits
[1484399125] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.71%) too high
[1484399185] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 246.62Mb (12.33%), Average OUT: 74.29Mb (3.71%)Total RX: 185794.12Gbits, Total TX: 130746.87Gbits
[1484399412] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.16%) too high
[1484399473] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 64.29Mb (64.29%), Average OUT: 12.70Mb (12.70%)Total RX: 28126.00Gbits, Total TX: 6541.28Gbits
[1484399545] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.63%) too high
[1484399593] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (79.63%) too high
[1484399605] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 262.21Mb (13.11%), Average OUT: 68.18Mb (3.41%)Total RX: 185906.70Gbits, Total TX: 130777.32Gbits
[1484399652] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (89.84%) too high
[1484399713] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (92.56%) too high
[1484399773] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (90.39%) too high
[1484399833] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (82.96%) too high
[1484399844] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671670.88%) too high
[1484399905] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 859.55Mb (42.98%), Average OUT: 70.43Mb (3.52%)Total RX: 186023.43Gbits, Total TX: 130831.75Gbits
[1484399953] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (96.08%) too high
[1484399965] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671673.08%) too high
[1484400025] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 290.24Mb (14.51%), Average OUT: 54.68Mb (2.73%)Total RX: 186025.70Gbits, Total TX: 130838.94Gbits
[1484400073] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (84.23%) too high
[1484400073] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1484400084] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.22%) too high
[1484400133] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (90.93%) too high
[1484400145] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 335.96Mb (16.80%), Average OUT: 79.75Mb (3.99%)Total RX: 186064.15Gbits, Total TX: 130813.19Gbits
[1484400624] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671661.09%) too high
[1484400685] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 336.44Mb (16.82%), Average OUT: 84.91Mb (4.25%)Total RX: 186234.45Gbits, Total TX: 130851.31Gbits
[1484400854] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (89.70%) too high
[1484400913] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 73.19Mb (73.19%), Average OUT: 16.05Mb (16.05%)Total RX: 28255.18Gbits, Total TX: 6561.59Gbits
[1484401033] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (82.41%) too high
[1484401045] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293809.66%) too high
[1484401093] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (75.21%) too high
[1484401105] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 244.89Mb (12.24%), Average OUT: 56.67Mb (2.83%)Total RX: 186363.44Gbits, Total TX: 130910.30Gbits
[1484401152] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (76.99%) too high
[1484401165] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.46%) too high
[1484401213] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;4;OK - Average IN: 74.37Mb (74.37%), Average OUT: 20.51Mb (20.51%)Total RX: 28278.20Gbits, Total TX: 6566.41Gbits
[1484401225] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 248.56Mb (12.43%), Average OUT: 630.77Mb (31.54%)Total RX: 186394.27Gbits, Total TX: 130917.83Gbits
[1484401273] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (79.69%) too high
[1484401333] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (76.19%) too high
[1484401393] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 70.17Mb (70.17%), Average OUT: 16.76Mb (16.76%)Total RX: 28291.77Gbits, Total TX: 6570.05Gbits
[1484401513] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.85%) too high
[1484401573] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 74.84Mb (74.84%), Average OUT: 15.58Mb (15.58%)Total RX: 28305.41Gbits, Total TX: 6573.06Gbits
[1484401705] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.07%) too high
[1484401753] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.86%) too high
[1484401765] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 264.64Mb (13.23%), Average OUT: 72.57Mb (3.63%)Total RX: 186528.45Gbits, Total TX: 130917.04Gbits
[1484401813] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (80.14%) too high
[1484401873] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (76.73%) too high
[1484401933] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (81.23%) too high
[1484402245] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.29%) too high
[1484402305] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 280.27Mb (14.01%), Average OUT: 109.35Mb (5.47%)Total RX: 186683.66Gbits, Total TX: 130954.32Gbits
[1484402353] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (90.80%) too high
[1484402413] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (81.40%) too high
[1484402473] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (91.62%) too high
[1484402533] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (79.88%) too high
[1484402725] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.14%) too high
[1484402785] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 249.27Mb (12.46%), Average OUT: 62.05Mb (3.10%)Total RX: 186807.28Gbits, Total TX: 130983.47Gbits
[1484402832] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (90.10%) too high
[1484402953] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (80.39%) too high
[1484403013] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (90.69%) too high
[1484403133] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (87.37%) too high
[1484403253] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.78%) too high
[1484403325] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.08%) too high
[1484403373] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (87.92%) too high
[1484403385] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 263.00Mb (13.15%), Average OUT: 87.39Mb (4.37%)Total RX: 186961.56Gbits, Total TX: 131023.65Gbits
[1484403433] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (91.56%) too high
[1484403493] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (87.70%) too high
[1484403745] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.42%) too high
[1484403805] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 254.62Mb (12.73%), Average OUT: 79.03Mb (3.95%)Total RX: 187075.40Gbits, Total TX: 131056.46Gbits
[1484403913] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.28%) too high
[1484403973] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (79.21%) too high
[1484404033] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 74.14Mb (74.14%), Average OUT: 13.69Mb (13.69%)Total RX: 28511.72Gbits, Total TX: 6613.05Gbits
[1484404045] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293809.16%) too high
[1484404105] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 271.36Mb (13.57%), Average OUT: 81.59Mb (4.08%)Total RX: 187154.56Gbits, Total TX: 131116.72Gbits
[1484404165] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.69%) too high
[1484404213] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (77.05%) too high
[1484404224] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 260.87Mb (13.04%), Average OUT: 664.59Mb (33.23%)Total RX: 187185.58Gbits, Total TX: 131125.88Gbits
[1484404273] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 59.49Mb (59.49%), Average OUT: 20.07Mb (20.07%)Total RX: 28528.09Gbits, Total TX: 6616.80Gbits
[1484404405] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671671.06%) too high
[1484404465] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 254.42Mb (12.72%), Average OUT: 65.83Mb (3.29%)Total RX: 187257.42Gbits, Total TX: 131147.71Gbits
[1484404525] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.33%) too high
[1484404585] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 898.44Mb (44.92%), Average OUT: 104.22Mb (5.21%)Total RX: 187328.60Gbits, Total TX: 131123.48Gbits
[1484404645] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293808.73%) too high
[1484404705] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 216.40Mb (10.82%), Average OUT: 80.72Mb (4.04%)Total RX: 187322.39Gbits, Total TX: 131134.61Gbits
[1484404753] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.09%) too high
[1484404811] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (83.02%) too high
[1484404872] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (84.37%) too high
[1484404933] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;4;OK - Average IN: 72.98Mb (72.98%), Average OUT: 13.10Mb (13.10%)Total RX: 28573.00Gbits, Total TX: 6629.19Gbits
[1484405366] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.05%) too high
[1484405413] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1484405424] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 261.72Mb (13.09%), Average OUT: 61.83Mb (3.09%)Total RX: 187505.14Gbits, Total TX: 131191.31Gbits
[1484405593] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (87.51%) too high
[1484405653] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 68.60Mb (68.60%), Average OUT: 8.28Mb (8.28%)Total RX: 28622.49Gbits, Total TX: 6638.26Gbits
[1484405845] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.04%) too high
[1484405892] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.66%) too high
[1484405906] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 251.80Mb (12.59%), Average OUT: 648.42Mb (32.42%)Total RX: 187623.66Gbits, Total TX: 131261.51Gbits
[1484405953] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 68.98Mb (68.98%), Average OUT: 14.50Mb (14.50%)Total RX: 28644.10Gbits, Total TX: 6641.28Gbits
[1484406373] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (79.58%) too high
[1484406385] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.67%) too high
[1484406433] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 65.34Mb (65.34%), Average OUT: 14.36Mb (14.36%)Total RX: 28675.28Gbits, Total TX: 6647.84Gbits
[1484406445] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 243.97Mb (12.20%), Average OUT: 77.02Mb (3.85%)Total RX: 187756.42Gbits, Total TX: 131262.96Gbits
[1484406805] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671661.16%) too high
[1484406866] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (120962256207.91%) too high
[1484406913] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.33%) too high
[1484406924] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 256.65Mb (12.83%), Average OUT: 79.30Mb (3.96%)Total RX: 187877.20Gbits, Total TX: 131300.62Gbits
[1484406973] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 67.14Mb (67.14%), Average OUT: 11.29Mb (11.29%)Total RX: 28708.87Gbits, Total TX: 6653.96Gbits
[1484407092] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.11%) too high
[1484407153] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [CHESSYS - F0/13];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10013 status: 2)
[1484407153] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 73.63Mb (73.63%), Average OUT: 14.65Mb (14.65%)Total RX: 28722.19Gbits, Total TX: 6656.40Gbits
[1484407192] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1484407213] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [CHESSYS - F0/13];OK;SOFT;2;OK - Average IN: 136.17Kb (0.14%), Average OUT: 2.52Mb (2.52%)Total RX: 95.98Gbits, Total TX: 828.21Gbits
[1484407213] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (90.61%) too high
[1484407257] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1484407273] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (91.24%) too high
[1484407318] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;3;(Host check timed out after 30.02 seconds)
[1484407333] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (79.26%) too high
[1484407376] HOST ALERT: INASOL-STEVEN;DOWN;HARD;4;PING CRITICAL - Packet loss = 100%
[1484407393] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (83.12%) too high
[1484407453] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (91.96%) too high
[1484407513] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (84.12%) too high
[1484407595] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484407653] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 473.92b (0.00%)Total RX: 17.62Kbits, Total TX: 507.17Mbits
[1484407656] HOST ALERT: INASOL-STEVEN;UP;HARD;4;PING OK - Packet loss = 0%, RTA = 9.03 ms
[1484407765] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.11%) too high
[1484407826] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 226.02Mb (11.30%), Average OUT: 54.33Mb (2.72%)Total RX: 188096.68Gbits, Total TX: 131363.01Gbits
[1484407878] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 67.02Mb (67.02%), Average OUT: 12.47Mb (12.47%)Total RX: 28781.47Gbits, Total TX: 6665.55Gbits
[1484407878] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1484408245] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.46%) too high
[1484408306] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 263.34Mb (13.17%), Average OUT: 69.31Mb (3.47%)Total RX: 188228.72Gbits, Total TX: 131396.94Gbits
[1484408415] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (84.40%) too high
[1484408432] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484408471] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 66.23Mb (66.23%), Average OUT: 10.94Mb (10.94%)Total RX: 28820.77Gbits, Total TX: 6672.20Gbits
[1484408484] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (127218924627.86%) too high
[1484408492] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 480.56b (0.00%)Total RX: 17.62Kbits, Total TX: 507.61Mbits
[1484408544] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 223.47Mb (11.17%), Average OUT: 650.74Mb (32.54%)Total RX: 188282.86Gbits, Total TX: 131448.00Gbits
[1484408726] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.53%) too high
[1484408774] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1484408785] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 240.68Mb (12.03%), Average OUT: 72.09Mb (3.60%)Total RX: 188335.87Gbits, Total TX: 131432.10Gbits
[1484408835] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.56%) too high
[1484408894] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 67.46Mb (67.46%), Average OUT: 10.15Mb (10.15%)Total RX: 28847.95Gbits, Total TX: 6676.95Gbits
[1484409204] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.15%) too high
[1484409265] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 267.03Mb (13.35%), Average OUT: 70.28Mb (3.51%)Total RX: 188442.34Gbits, Total TX: 131468.97Gbits
[1484409272] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484409334] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 470.72b (0.00%)Total RX: 17.62Kbits, Total TX: 508.04Mbits
[1484409565] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671668.00%) too high
[1484409626] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 230.62Mb (11.53%), Average OUT: 625.40Mb (31.27%)Total RX: 188524.04Gbits, Total TX: 131529.00Gbits
[1484410166] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.93%) too high
[1484410226] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 237.33Mb (11.87%), Average OUT: 80.46Mb (4.02%)Total RX: 188670.18Gbits, Total TX: 131538.05Gbits
[1484410525] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671661.95%) too high
[1484410584] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 231.40Mb (11.57%), Average OUT: 104.15Mb (5.21%)Total RX: 188754.86Gbits, Total TX: 131569.92Gbits
[1484410705] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671670.67%) too high
[1484410765] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 257.17Mb (12.86%), Average OUT: 86.49Mb (4.32%)Total RX: 188803.43Gbits, Total TX: 131586.17Gbits
[1484410946] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256198.08%) too high
[1484411006] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 237.34Mb (11.87%), Average OUT: 676.74Mb (33.84%)Total RX: 188859.84Gbits, Total TX: 131644.58Gbits
[1484411125] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293808.12%) too high
[1484411183] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.45, 2.10, 2.97
[1484411186] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 227.44Mb (11.37%), Average OUT: 92.98Mb (4.65%)Total RX: 188903.65Gbits, Total TX: 131660.77Gbits
[1484411306] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.51%) too high
[1484411365] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 248.10Mb (12.41%), Average OUT: 61.72Mb (3.09%)Total RX: 188947.42Gbits, Total TX: 131639.74Gbits
[1484411846] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.25%) too high
[1484411904] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 3.26, 3.16, 3.01
[1484411905] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 216.13Mb (10.81%), Average OUT: 58.95Mb (2.95%)Total RX: 189069.91Gbits, Total TX: 131673.58Gbits
[1484411963] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 2.34, 3.01, 2.98
[1484412023] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 3.58, 3.31, 3.08
[1484412023] SERVICE FLAPPING ALERT: OS-MONITOREO;Carga OS-MONITOREO;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1484412386] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.11%) too high
[1484412446] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 237.28Mb (11.86%), Average OUT: 59.27Mb (2.96%)Total RX: 189196.61Gbits, Total TX: 131705.26Gbits
[1484412985] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.52%) too high
[1484413046] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 180.77Mb (9.04%), Average OUT: 76.17Mb (3.81%)Total RX: 189314.70Gbits, Total TX: 131743.08Gbits
[1484413163] SERVICE FLAPPING ALERT: OS-MONITOREO;Carga OS-MONITOREO;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1484413344] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671667.78%) too high
[1484413405] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (120962256196.59%) too high
[1484413466] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 726.15Mb (36.31%), Average OUT: 72.53Mb (3.63%)Total RX: 189424.44Gbits, Total TX: 131775.19Gbits
[1484413526] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293803.16%) too high
[1484413586] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 156.75Mb (7.84%), Average OUT: 100.14Mb (5.01%)Total RX: 189407.96Gbits, Total TX: 131786.32Gbits
[1484413766] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293803.21%) too high
[1484413826] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 146.77Mb (7.34%), Average OUT: 645.76Mb (32.29%)Total RX: 189442.55Gbits, Total TX: 131837.71Gbits
[1484413886] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.30%) too high
[1484413946] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 131.05Mb (6.55%), Average OUT: 55.34Mb (2.77%)Total RX: 189458.20Gbits, Total TX: 131810.51Gbits
[1484414246] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256201.62%) too high
[1484414305] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 161.11Mb (8.06%), Average OUT: 64.76Mb (3.24%)Total RX: 189513.16Gbits, Total TX: 131869.66Gbits
[1484414366] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.08%) too high
[1484414426] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 153.85Mb (7.69%), Average OUT: 79.65Mb (3.98%)Total RX: 189533.27Gbits, Total TX: 131843.98Gbits
[1484414653] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [CHESSYS - F0/13];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10013 status: 2)
[1484414713] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [CHESSYS - F0/13];OK;SOFT;2;OK - Average IN: 17.06Kb (0.02%), Average OUT: 703.07Kb (0.70%)Total RX: 96.93Gbits, Total TX: 846.04Gbits
[1484414905] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.26%) too high
[1484414966] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 133.80Mb (6.69%), Average OUT: 94.34Mb (4.72%)Total RX: 189608.97Gbits, Total TX: 131882.77Gbits
[1484415266] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.27%) too high
[1484415326] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 147.04Mb (7.35%), Average OUT: 61.27Mb (3.06%)Total RX: 189658.10Gbits, Total TX: 131913.49Gbits
[1484415446] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (119011252067.39%) too high
[1484415506] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 150.74Mb (7.54%), Average OUT: 56.91Mb (2.85%)Total RX: 189682.47Gbits, Total TX: 131922.91Gbits
[1484415684] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (127218924624.05%) too high
[1484415694] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484415746] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 159.50Mb (7.97%), Average OUT: 608.16Mb (30.41%)Total RX: 189718.07Gbits, Total TX: 131969.66Gbits
[1484415755] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 475.68b (0.00%)Total RX: 17.62Kbits, Total TX: 511.38Mbits
[1484415865] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.12%) too high
[1484415925] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 144.99Mb (7.25%), Average OUT: 638.88Mb (31.94%)Total RX: 189744.29Gbits, Total TX: 131980.70Gbits
[1484416463] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.91, 2.45, 2.93
[1484416524] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.85%) too high
[1484416586] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 140.42Mb (7.02%), Average OUT: 65.94Mb (3.30%)Total RX: 189838.19Gbits, Total TX: 131982.07Gbits
[1484416763] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 5.45, 3.18, 3.04
[1484417065] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.10%) too high
[1484417124] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 159.29Mb (7.96%), Average OUT: 58.30Mb (2.91%)Total RX: 189913.16Gbits, Total TX: 132014.25Gbits
[1484417605] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.37%) too high
[1484417666] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 146.62Mb (7.33%), Average OUT: 630.68Mb (31.53%)Total RX: 189999.06Gbits, Total TX: 132083.58Gbits
[1484418144] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 2.30, 2.69, 2.99
[1484418203] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 3.63, 2.86, 3.02
[1484418205] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.51%) too high
[1484418265] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 111.13Mb (5.56%), Average OUT: 69.37Mb (3.47%)Total RX: 190075.23Gbits, Total TX: 132085.83Gbits
[1484418744] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.93%) too high
[1484418804] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 116.84Mb (5.84%), Average OUT: 58.13Mb (2.91%)Total RX: 190145.48Gbits, Total TX: 132118.93Gbits
[1484418951] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1484419018] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1484419078] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1484419137] HOST ALERT: INASOL-STEVEN;DOWN;HARD;4;PING CRITICAL - Packet loss = 100%
[1484419285] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.65%) too high
[1484419344] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 131.02Mb (6.55%), Average OUT: 61.23Mb (3.06%)Total RX: 190212.66Gbits, Total TX: 132151.69Gbits
[1484419825] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.17%) too high
[1484419885] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 105.08Mb (5.25%), Average OUT: 48.98Mb (2.45%)Total RX: 190277.06Gbits, Total TX: 132185.56Gbits
[1484420006] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256199.43%) too high
[1484420065] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 130.76Mb (6.54%), Average OUT: 647.29Mb (32.36%)Total RX: 190299.75Gbits, Total TX: 132229.88Gbits
[1484420431] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484420485] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.88%) too high
[1484420490] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 459.44b (0.00%)Total RX: 17.62Kbits, Total TX: 513.83Mbits
[1484420546] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 128.97Mb (6.45%), Average OUT: 65.74Mb (3.29%)Total RX: 190357.34Gbits, Total TX: 132223.15Gbits
[1484420723] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 2.35, 2.69, 2.99
[1484420843] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 4.43, 2.99, 3.04
[1484420845] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293802.10%) too high
[1484420905] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 108.60Mb (5.43%), Average OUT: 51.00Mb (2.55%)Total RX: 190400.17Gbits, Total TX: 132241.69Gbits
[1484421144] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.79%) too high
[1484421206] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 118.58Mb (5.93%), Average OUT: 48.95Mb (2.45%)Total RX: 190439.21Gbits, Total TX: 132256.80Gbits
[1484421686] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (119011252064.73%) too high
[1484421692] HOST ALERT: INASOL-STEVEN;UP;HARD;4;PING OK - Packet loss = 0%, RTA = 7.20 ms
[1484421746] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 125.22Mb (6.26%), Average OUT: 86.25Mb (4.31%)Total RX: 190502.90Gbits, Total TX: 132293.45Gbits
[1484422104] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.41%) too high
[1484422165] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 76.82Mb (3.84%), Average OUT: 55.62Mb (2.78%)Total RX: 190542.80Gbits, Total TX: 132324.86Gbits
[1484422405] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671661.48%) too high
[1484422464] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 83.37Mb (4.17%), Average OUT: 33.06Mb (1.65%)Total RX: 190568.48Gbits, Total TX: 132339.47Gbits
[1484423364] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (127218924620.23%) too high
[1484423424] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 662.77Mb (33.14%), Average OUT: 67.08Mb (3.35%)Total RX: 190695.69Gbits, Total TX: 132394.66Gbits
[1484423485] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256196.97%) too high
[1484423545] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 82.04Mb (4.10%), Average OUT: 55.22Mb (2.76%)Total RX: 190671.24Gbits, Total TX: 132437.15Gbits
[1484423906] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.26%) too high
[1484423964] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 89.60Mb (4.48%), Average OUT: 76.54Mb (3.83%)Total RX: 190708.06Gbits, Total TX: 132429.78Gbits
[1484424329] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484424389] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 476.24b (0.00%)Total RX: 17.62Kbits, Total TX: 515.85Mbits
[1484424564] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (127218924618.67%) too high
[1484424626] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 119.90Mb (5.99%), Average OUT: 57.89Mb (2.89%)Total RX: 190769.75Gbits, Total TX: 132460.91Gbits
[1484424926] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256199.30%) too high
[1484424985] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 107.74Mb (5.39%), Average OUT: 645.56Mb (32.28%)Total RX: 190813.48Gbits, Total TX: 132520.96Gbits
[1484425526] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.53%) too high
[1484425584] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 96.89Mb (4.84%), Average OUT: 62.12Mb (3.11%)Total RX: 190875.24Gbits, Total TX: 132530.46Gbits
[1484425626] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1484425678] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1484425738] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1484425798] HOST ALERT: INASOL-STEVEN;DOWN;HARD;4;(Host check timed out after 30.01 seconds)
[1484426064] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (127218924619.30%) too high
[1484426126] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 82.77Mb (4.14%), Average OUT: 58.15Mb (2.91%)Total RX: 190920.16Gbits, Total TX: 132565.22Gbits
[1484426604] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.37%) too high
[1484426665] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 75.76Mb (3.79%), Average OUT: 47.01Mb (2.35%)Total RX: 190963.43Gbits, Total TX: 132597.38Gbits
[1484427206] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (117122184571.03%) too high
[1484427265] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 65.65Mb (3.28%), Average OUT: 57.55Mb (2.88%)Total RX: 191008.36Gbits, Total TX: 132636.17Gbits
[1484427389] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484427448] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 459.44b (0.00%)Total RX: 18.91Kbits, Total TX: 517.44Mbits
[1484427745] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671660.35%) too high
[1484427804] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 82.02Mb (4.10%), Average OUT: 59.71Mb (2.99%)Total RX: 191048.39Gbits, Total TX: 132695.21Gbits
[1484427865] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.57%) too high
[1484427925] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 72.18Mb (3.61%), Average OUT: 68.89Mb (3.44%)Total RX: 191057.78Gbits, Total TX: 132668.26Gbits
[1484429185] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.46%) too high
[1484429245] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 89.93Mb (4.50%), Average OUT: 55.55Mb (2.78%)Total RX: 191165.96Gbits, Total TX: 132736.68Gbits
[1484429844] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.48%) too high
[1484429905] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 76.45Mb (3.82%), Average OUT: 43.24Mb (2.16%)Total RX: 191218.16Gbits, Total TX: 132769.97Gbits
[1484430326] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256196.73%) too high
[1484430383] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 82.93Mb (4.15%), Average OUT: 48.60Mb (2.43%)Total RX: 191255.45Gbits, Total TX: 132833.42Gbits
[1484430446] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (117122184570.04%) too high
[1484430448] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484430504] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 68.75Mb (3.44%), Average OUT: 46.70Mb (2.34%)Total RX: 191263.98Gbits, Total TX: 132804.71Gbits
[1484430509] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 469.52b (0.00%)Total RX: 18.91Kbits, Total TX: 519.03Mbits
[1484430805] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (119011252064.25%) too high
[1484430866] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 68.60Mb (3.43%), Average OUT: 47.93Mb (2.40%)Total RX: 191288.80Gbits, Total TX: 132821.41Gbits
[1484431166] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.04%) too high
[1484431225] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 661.55Mb (33.08%), Average OUT: 687.76Mb (34.39%)Total RX: 191349.04Gbits, Total TX: 132878.23Gbits
[1484431286] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256196.66%) too high
[1484431345] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 77.53Mb (3.88%), Average OUT: 60.92Mb (3.05%)Total RX: 191323.86Gbits, Total TX: 132887.52Gbits
[1484431585] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.73%) too high
[1484431644] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 71.29Mb (3.56%), Average OUT: 55.74Mb (2.79%)Total RX: 191344.01Gbits, Total TX: 132875.60Gbits
[1484431766] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256196.20%) too high
[1484431824] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 73.57Mb (3.68%), Average OUT: 48.45Mb (2.42%)Total RX: 191356.85Gbits, Total TX: 132884.51Gbits
[1484432246] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256195.83%) too high
[1484432304] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (127218924618.95%) too high
[1484432366] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 52.02Mb (2.60%), Average OUT: 38.02Mb (1.90%)Total RX: 191393.78Gbits, Total TX: 132908.03Gbits
[1484433324] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (127218924619.97%) too high
[1484433386] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 75.87Mb (3.79%), Average OUT: 75.02Mb (3.75%)Total RX: 191460.60Gbits, Total TX: 133001.56Gbits
[1484433505] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.94%) too high
[1484433566] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 70.60Mb (3.53%), Average OUT: 63.89Mb (3.19%)Total RX: 191472.50Gbits, Total TX: 132980.21Gbits
[1484433986] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.84%) too high
[1484434045] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 55.71Mb (2.79%), Average OUT: 62.81Mb (3.14%)Total RX: 191501.11Gbits, Total TX: 133010.33Gbits
[1484434646] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (117122184570.79%) too high
[1484434705] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 61.07Mb (3.05%), Average OUT: 53.92Mb (2.70%)Total RX: 191539.53Gbits, Total TX: 133047.92Gbits
[1484435184] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (127218924619.30%) too high
[1484435246] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 59.82Mb (2.99%), Average OUT: 44.72Mb (2.24%)Total RX: 191571.87Gbits, Total TX: 133078.96Gbits
[1484435728] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484435788] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 483.68b (0.00%)Total RX: 18.91Kbits, Total TX: 521.77Mbits
[1484436266] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (119011252064.39%) too high
[1484436324] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 63.82Mb (3.19%), Average OUT: 69.54Mb (3.48%)Total RX: 191635.17Gbits, Total TX: 133151.60Gbits
[1484436746] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (117122184571.24%) too high
[1484436804] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 45.47Mb (2.27%), Average OUT: 88.77Mb (4.44%)Total RX: 191661.35Gbits, Total TX: 133186.79Gbits
[1484437225] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.68%) too high
[1484437286] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 64.05Mb (3.20%), Average OUT: 78.09Mb (3.90%)Total RX: 191687.88Gbits, Total TX: 133221.10Gbits
[1484437948] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484438008] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 480.56b (0.00%)Total RX: 18.91Kbits, Total TX: 522.92Mbits
[1484438306] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.78%) too high
[1484438365] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 67.08Mb (3.35%), Average OUT: 56.32Mb (2.82%)Total RX: 191753.09Gbits, Total TX: 133287.44Gbits
[1484438846] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.87%) too high
[1484438904] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 50.03Mb (2.50%), Average OUT: 74.10Mb (3.70%)Total RX: 191780.85Gbits, Total TX: 133322.31Gbits
[1484439266] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256195.05%) too high
[1484439325] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (125062671660.00%) too high
[1484439385] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 44.69Mb (2.23%), Average OUT: 72.22Mb (3.61%)Total RX: 191803.62Gbits, Total TX: 133355.30Gbits
[1484439806] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.75%) too high
[1484439866] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 37.77Mb (1.89%), Average OUT: 63.16Mb (3.16%)Total RX: 191823.28Gbits, Total TX: 133389.25Gbits
[1484440346] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.08%) too high
[1484440404] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 42.60Mb (2.13%), Average OUT: 59.09Mb (2.95%)Total RX: 191845.41Gbits, Total TX: 133424.37Gbits
[1484440592] HOST ALERT: INASOL-STEVEN;UP;HARD;4;PING OK - Packet loss = 0%, RTA = 8.32 ms
[1484440945] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.88%) too high
[1484441006] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 40.19Mb (2.01%), Average OUT: 50.59Mb (2.53%)Total RX: 191869.01Gbits, Total TX: 133456.77Gbits
[1484441606] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (119011252063.54%) too high
[1484441665] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 41.42Mb (2.07%), Average OUT: 634.78Mb (31.74%)Total RX: 191899.67Gbits, Total TX: 133527.53Gbits
[1484442266] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.66%) too high
[1484442326] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 37.09Mb (1.85%), Average OUT: 55.43Mb (2.77%)Total RX: 191924.15Gbits, Total TX: 133526.37Gbits
[1484442926] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.26%) too high
[1484442986] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 37.79Mb (1.89%), Average OUT: 38.78Mb (1.94%)Total RX: 191953.40Gbits, Total TX: 133561.73Gbits
[1484443466] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293797.52%) too high
[1484443525] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 27.74Mb (1.39%), Average OUT: 50.56Mb (2.53%)Total RX: 191970.83Gbits, Total TX: 133623.80Gbits
[1484443586] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256194.91%) too high
[1484443646] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 31.23Mb (1.56%), Average OUT: 54.22Mb (2.71%)Total RX: 191974.66Gbits, Total TX: 133595.16Gbits
[1484444246] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.66%) too high
[1484444306] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 31.21Mb (1.56%), Average OUT: 37.48Mb (1.87%)Total RX: 192002.84Gbits, Total TX: 133629.56Gbits
[1484444966] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.91%) too high
[1484445025] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 29.91Mb (1.50%), Average OUT: 49.82Mb (2.49%)Total RX: 192029.72Gbits, Total TX: 133664.56Gbits
[1484445506] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.46%) too high
[1484445566] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 27.46Mb (1.37%), Average OUT: 632.49Mb (31.62%)Total RX: 192047.25Gbits, Total TX: 133732.24Gbits
[1484446106] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (119011252063.87%) too high
[1484446166] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 26.05Mb (1.30%), Average OUT: 62.52Mb (3.13%)Total RX: 192065.02Gbits, Total TX: 133735.19Gbits
[1484446586] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.31%) too high
[1484446646] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 24.74Mb (1.24%), Average OUT: 74.85Mb (3.74%)Total RX: 192077.78Gbits, Total TX: 133769.01Gbits
[1484447126] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.22%) too high
[1484447186] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 34.11Mb (1.71%), Average OUT: 72.71Mb (3.64%)Total RX: 192093.80Gbits, Total TX: 133803.16Gbits
[1484447606] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293798.24%) too high
[1484447666] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 47.90Mb (2.40%), Average OUT: 64.39Mb (3.22%)Total RX: 192110.67Gbits, Total TX: 133836.21Gbits
[1484448086] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293801.93%) too high
[1484448146] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 34.37Mb (1.72%), Average OUT: 127.23Mb (6.36%)Total RX: 192126.98Gbits, Total TX: 133875.16Gbits
[1484448506] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.95%) too high
[1484448566] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 599.09Mb (29.95%), Average OUT: 83.13Mb (4.16%)Total RX: 192173.97Gbits, Total TX: 133907.38Gbits
[1484448626] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293797.44%) too high
[1484448686] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 28.60Mb (1.43%), Average OUT: 647.34Mb (32.37%)Total RX: 192142.94Gbits, Total TX: 133950.91Gbits
[1484448986] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.10%) too high
[1484449046] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 28.75Mb (1.44%), Average OUT: 95.81Mb (4.79%)Total RX: 192152.37Gbits, Total TX: 133942.95Gbits
[1484449406] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.34%) too high
[1484449466] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 26.89Mb (1.34%), Average OUT: 64.19Mb (3.21%)Total RX: 192164.08Gbits, Total TX: 133975.83Gbits
[1484449886] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293797.32%) too high
[1484449946] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (125062671659.68%) too high
[1484450005] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 30.48Mb (1.52%), Average OUT: 54.52Mb (2.73%)Total RX: 192178.14Gbits, Total TX: 134007.49Gbits
[1484450546] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.78%) too high
[1484450606] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 32.80Mb (1.64%), Average OUT: 58.54Mb (2.93%)Total RX: 192196.76Gbits, Total TX: 134042.23Gbits
[1484451086] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.55%) too high
[1484451146] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 49.95Mb (2.50%), Average OUT: 632.35Mb (31.62%)Total RX: 192217.38Gbits, Total TX: 134113.55Gbits
[1484451686] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.73%) too high
[1484451746] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 41.17Mb (2.06%), Average OUT: 43.74Mb (2.19%)Total RX: 192243.72Gbits, Total TX: 134108.76Gbits
[1484452406] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.41%) too high
[1484452466] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 27.08Mb (1.35%), Average OUT: 40.04Mb (2.00%)Total RX: 192268.38Gbits, Total TX: 134144.41Gbits
[1484453127] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.38%) too high
[1484453186] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 32.78Mb (1.64%), Average OUT: 56.12Mb (2.81%)Total RX: 192287.34Gbits, Total TX: 134181.45Gbits
[1484453786] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.11%) too high
[1484453817] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1484453847] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 29.65Mb (1.48%), Average OUT: 62.52Mb (3.13%)Total RX: 192305.37Gbits, Total TX: 134216.09Gbits
[1484453878] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;2;(Host check timed out after 30.02 seconds)
[1484453937] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1484453998] HOST ALERT: INASOL-STEVEN;DOWN;HARD;4;(Host check timed out after 30.01 seconds)
[1484454146] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293797.34%) too high
[1484454206] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 25.31Mb (1.27%), Average OUT: 47.53Mb (2.38%)Total RX: 192314.22Gbits, Total TX: 134233.35Gbits
[1484454446] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.15%) too high
[1484454506] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 25.16Mb (1.26%), Average OUT: 61.82Mb (3.09%)Total RX: 192321.76Gbits, Total TX: 134250.75Gbits
[1484454986] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.84%) too high
[1484455046] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 25.49Mb (1.27%), Average OUT: 39.71Mb (1.99%)Total RX: 192336.49Gbits, Total TX: 134280.92Gbits
[1484455706] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.29%) too high
[1484455766] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 25.77Mb (1.29%), Average OUT: 59.15Mb (2.96%)Total RX: 192357.74Gbits, Total TX: 134316.99Gbits
[1484455993] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [CHESSYS - F0/13];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10013 status: 2)
[1484456012] HOST ALERT: INASOL-STEVEN;UP;HARD;4;PING OK - Packet loss = 0%, RTA = 8.21 ms
[1484456053] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [CHESSYS - F0/13];OK;SOFT;2;OK - Average IN: 6.08Kb (0.01%), Average OUT: 12.18Kb (0.01%)Total RX: 97.24Gbits, Total TX: 849.92Gbits
[1484456366] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.17%) too high