[1484197353] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;1;WARNING - load average: 5.30, 2.81, 2.28
[1484197413] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;OK;SOFT;2;OK - load average: 3.26, 2.59, 2.23
[1484197473] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;SOFT;1;CRITICAL - load average: 6.20, 3.22, 2.45
[1484197533] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;SOFT;2;CRITICAL - load average: 6.19, 3.47, 2.57
[1484197593] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;3;WARNING - load average: 5.56, 3.55, 2.64
[1484197594] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.97%) too high
[1484197653] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 26.46Mb (1.32%), Average OUT: 57.69Mb (2.88%)Total RX: 147255.36Gbits, Total TX: 116930.13Gbits
[1484197653] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 4.69, 3.49, 2.67
[1484197893] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;OK;HARD;4;OK - load average: 1.88, 2.48, 2.43
[1484198074] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293797.37%) too high
[1484198133] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 25.97Mb (1.30%), Average OUT: 49.08Mb (2.45%)Total RX: 147268.65Gbits, Total TX: 116989.62Gbits
[1484198253] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.94%) too high
[1484198313] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 28.30Mb (1.42%), Average OUT: 64.49Mb (3.22%)Total RX: 147273.34Gbits, Total TX: 116966.28Gbits
[1484198732] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484198790] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 465.36b (0.00%)Total RX: 8.99Kbits, Total TX: 398.56Mbits
[1484198853] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.73%) too high
[1484198913] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 38.83Mb (1.94%), Average OUT: 54.82Mb (2.74%)Total RX: 147290.71Gbits, Total TX: 116998.79Gbits
[1484199093] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;1;WARNING - load average: 4.37, 2.88, 2.51
[1484199154] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;2;WARNING - load average: 4.67, 3.06, 2.58
[1484199212] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671659.03%) too high
[1484199213] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;3;WARNING - load average: 4.28, 3.07, 2.61
[1484199272] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 37.03Mb (1.85%), Average OUT: 61.73Mb (3.09%)Total RX: 147303.53Gbits, Total TX: 117054.35Gbits
[1484199273] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 4.42, 3.16, 2.66
[1484199453] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.82%) too high
[1484199513] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 32.88Mb (1.64%), Average OUT: 51.99Mb (2.60%)Total RX: 147311.44Gbits, Total TX: 117033.84Gbits
[1484200168] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (134158138690.16%) too high
[1484200231] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 45.92Mb (2.30%), Average OUT: 55.48Mb (2.77%)Total RX: 147340.20Gbits, Total TX: 117069.27Gbits
[1484200473] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 4.99, 3.57, 3.07
[1484200832] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.02%) too high
[1484200892] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 57.52Mb (2.88%), Average OUT: 53.18Mb (2.66%)Total RX: 147372.75Gbits, Total TX: 117103.25Gbits
[1484200952] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484201012] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 472b (0.00%)Total RX: 8.99Kbits, Total TX: 399.71Mbits
[1484201372] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.86%) too high
[1484201433] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 621.86Mb (31.09%), Average OUT: 636.67Mb (31.83%)Total RX: 147435.70Gbits, Total TX: 117171.93Gbits
[1484201489] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (131762457641.13%) too high
[1484201552] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 50.87Mb (2.54%), Average OUT: 73.48Mb (3.67%)Total RX: 147407.29Gbits, Total TX: 117180.20Gbits
[1484201912] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.88%) too high
[1484201973] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 63.06Mb (3.15%), Average OUT: 49.74Mb (2.49%)Total RX: 147429.81Gbits, Total TX: 117171.97Gbits
[1484202148] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671661.11%) too high
[1484202213] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 64.89Mb (3.24%), Average OUT: 55.08Mb (2.75%)Total RX: 147445.96Gbits, Total TX: 117218.43Gbits
[1484202513] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.52%) too high
[1484202570] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 667.37Mb (33.37%), Average OUT: 53.21Mb (2.66%)Total RX: 147504.67Gbits, Total TX: 117204.63Gbits
[1484202629] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671660.70%) too high
[1484202692] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 74.78Mb (3.74%), Average OUT: 58.87Mb (2.94%)Total RX: 147479.33Gbits, Total TX: 117211.28Gbits
[1484203171] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (119011252063.97%) too high
[1484203233] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 97.46Mb (4.87%), Average OUT: 57.87Mb (2.89%)Total RX: 147528.15Gbits, Total TX: 117242.12Gbits
[1484204253] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.43%) too high
[1484204310] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 166.30Mb (8.32%), Average OUT: 85.88Mb (4.29%)Total RX: 147667.49Gbits, Total TX: 117311.90Gbits
[1484204432] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (115292150444.30%) too high
[1484204492] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 220.20Mb (11.01%), Average OUT: 75.81Mb (3.79%)Total RX: 147703.47Gbits, Total TX: 117361.08Gbits
[1484204672] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (117122184571.40%) too high
[1484204732] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 220.15Mb (11.01%), Average OUT: 87.25Mb (4.36%)Total RX: 147754.47Gbits, Total TX: 117346.29Gbits
[1484204909] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (127218924626.38%) too high
[1484204972] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 236.94Mb (11.85%), Average OUT: 61.45Mb (3.07%)Total RX: 147805.09Gbits, Total TX: 117362.94Gbits
[1484205149] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (131762457641.65%) too high
[1484205209] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (122978293806.06%) too high
[1484205272] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 185.40Mb (9.27%), Average OUT: 63.94Mb (3.20%)Total RX: 147872.43Gbits, Total TX: 117382.85Gbits
[1484205392] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293804.76%) too high
[1484205452] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 207.71Mb (10.39%), Average OUT: 56.34Mb (2.82%)Total RX: 147906.86Gbits, Total TX: 117392.84Gbits
[1484205688] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (134158138690.28%) too high
[1484205753] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 261.40Mb (13.07%), Average OUT: 77.67Mb (3.88%)Total RX: 147979.91Gbits, Total TX: 117414.34Gbits
[1484206172] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.41%) too high
[1484206232] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 269.96Mb (13.50%), Average OUT: 59.09Mb (2.95%)Total RX: 148095.78Gbits, Total TX: 117445.85Gbits
[1484206713] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.32%) too high
[1484206769] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 263.80Mb (13.19%), Average OUT: 62.05Mb (3.10%)Total RX: 148224.76Gbits, Total TX: 117481.72Gbits
[1484207192] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.12%) too high
[1484207252] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 298.50Mb (14.92%), Average OUT: 71.94Mb (3.60%)Total RX: 148372.95Gbits, Total TX: 117514.98Gbits
[1484207492] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293814.61%) too high
[1484207553] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 368.68Mb (18.43%), Average OUT: 73.31Mb (3.67%)Total RX: 148478.21Gbits, Total TX: 117571.10Gbits
[1484207669] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.54%) too high
[1484207730] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 413.37Mb (20.67%), Average OUT: 71.24Mb (3.56%)Total RX: 148549.65Gbits, Total TX: 117549.06Gbits
[1484207852] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293811.74%) too high
[1484207912] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 833.42Mb (41.67%), Average OUT: 93.41Mb (4.67%)Total RX: 148640.11Gbits, Total TX: 117565.47Gbits
[1484207972] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293812.56%) too high
[1484208032] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 382.01Mb (19.10%), Average OUT: 66.16Mb (3.31%)Total RX: 148648.42Gbits, Total TX: 117573.98Gbits
[1484208199] SERVICE ALERT: INASOL-BGP01;INASOL [INTERNET REDCA - G0/0/0];WARNING;SOFT;1;WARNING - OUT bandwidth (83.66%) too high
[1484208258] SERVICE ALERT: INASOL-BGP01;INASOL [INTERNET REDCA - G0/0/0];OK;SOFT;2;OK - Average IN: 12.26Mb (12.26%), Average OUT: 19.17Mb (19.17%)Total RX: 22203.03Gbits, Total TX: 29396.43Gbits
[1484208452] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484208512] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 477.36b (0.00%)Total RX: 8.99Kbits, Total TX: 403.61Mbits
[1484208512] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293801.54%) too high
[1484208573] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 273.12Mb (13.66%), Average OUT: 96.21Mb (4.81%)Total RX: 148816.66Gbits, Total TX: 117620.46Gbits
[1484208872] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (119011252073.35%) too high
[1484208932] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 297.06Mb (14.85%), Average OUT: 109.40Mb (5.47%)Total RX: 148912.01Gbits, Total TX: 117692.57Gbits
[1484209113] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256208.43%) too high
[1484209171] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 872.66Mb (43.63%), Average OUT: 96.66Mb (4.83%)Total RX: 149017.39Gbits, Total TX: 117716.13Gbits
[1484209229] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (127218924631.88%) too high
[1484209290] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 306.22Mb (15.31%), Average OUT: 89.13Mb (4.46%)Total RX: 149019.38Gbits, Total TX: 117692.31Gbits
[1484209653] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (113518425046.00%) too high
[1484209708] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 359.78Mb (17.99%), Average OUT: 72.65Mb (3.63%)Total RX: 149160.61Gbits, Total TX: 117723.24Gbits
[1484210072] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.61%) too high
[1484210133] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 364.61Mb (18.23%), Average OUT: 61.93Mb (3.10%)Total RX: 149307.38Gbits, Total TX: 117755.91Gbits
[1484210252] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293813.39%) too high
[1484210312] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 314.85Mb (15.74%), Average OUT: 65.56Mb (3.28%)Total RX: 149366.76Gbits, Total TX: 117800.75Gbits
[1484210612] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (115292150437.02%) too high
[1484210672] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 983.72Mb (49.19%), Average OUT: 67.45Mb (3.37%)Total RX: 149531.87Gbits, Total TX: 117791.38Gbits
[1484210733] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256214.44%) too high
[1484210789] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 361.84Mb (18.09%), Average OUT: 79.31Mb (3.97%)Total RX: 149544.08Gbits, Total TX: 117800.52Gbits
[1484210913] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (113518425059.12%) too high
[1484210969] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 316.99Mb (15.85%), Average OUT: 71.79Mb (3.59%)Total RX: 149605.27Gbits, Total TX: 117813.83Gbits
[1484211077] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.79%) too high
[1484211139] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 69.22Mb (69.22%), Average OUT: 17.43Mb (17.43%)Total RX: 20194.68Gbits, Total TX: 5042.75Gbits
[1484211511] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.46%) too high
[1484211569] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 383.27Mb (19.16%), Average OUT: 84.25Mb (4.21%)Total RX: 149812.12Gbits, Total TX: 117859.95Gbits
[1484211932] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.59%) too high
[1484211992] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 322.79Mb (16.14%), Average OUT: 76.93Mb (3.85%)Total RX: 149966.45Gbits, Total TX: 117893.05Gbits
[1484212172] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671677.00%) too high
[1484212232] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 378.29Mb (18.91%), Average OUT: 105.91Mb (5.30%)Total RX: 150060.15Gbits, Total TX: 117945.95Gbits
[1484212352] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293815.61%) too high
[1484212412] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 410.88Mb (20.54%), Average OUT: 53.53Mb (2.68%)Total RX: 150131.29Gbits, Total TX: 117956.73Gbits
[1484212472] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.29%) too high
[1484212532] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 402.81Mb (20.14%), Average OUT: 86.25Mb (4.31%)Total RX: 150180.10Gbits, Total TX: 117932.57Gbits
[1484212773] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256211.11%) too high
[1484212828] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 392.19Mb (19.61%), Average OUT: 80.73Mb (4.04%)Total RX: 150298.17Gbits, Total TX: 117988.12Gbits
[1484212892] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (115292150437.37%) too high
[1484212953] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 419.29Mb (20.96%), Average OUT: 59.53Mb (2.98%)Total RX: 150349.87Gbits, Total TX: 117961.91Gbits
[1484213192] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (127218924640.48%) too high
[1484213252] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 442.48Mb (22.12%), Average OUT: 635.56Mb (31.78%)Total RX: 150480.90Gbits, Total TX: 118018.97Gbits
[1484213369] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (131762457642.16%) too high
[1484213432] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 392.47Mb (19.62%), Average OUT: 68.02Mb (3.40%)Total RX: 150551.19Gbits, Total TX: 117996.50Gbits
[1484213479] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (77.78%) too high
[1484213537] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 63.76Mb (63.76%), Average OUT: 22.15Mb (22.15%)Total RX: 20338.16Gbits, Total TX: 5074.82Gbits
[1484213610] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256213.61%) too high
[1484213672] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 400.63Mb (20.03%), Average OUT: 72.03Mb (3.60%)Total RX: 150647.25Gbits, Total TX: 118049.01Gbits
[1484213719] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.64%) too high
[1484213732] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484213777] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 67.23Mb (67.23%), Average OUT: 18.30Mb (18.30%)Total RX: 20354.99Gbits, Total TX: 5079.18Gbits
[1484213792] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 470b (0.00%)Total RX: 8.99Kbits, Total TX: 406.36Mbits
[1484213849] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.19%) too high
[1484213912] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 424.39Mb (21.22%), Average OUT: 74.06Mb (3.70%)Total RX: 150751.08Gbits, Total TX: 118032.85Gbits
[1484214329] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (131762457642.44%) too high
[1484214390] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 490.00Mb (24.50%), Average OUT: 70.65Mb (3.53%)Total RX: 150957.21Gbits, Total TX: 118068.58Gbits
[1484214730] HOST ALERT: INASOL-CyM_ABOGADOS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1484214749] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (131762457642.32%) too high
[1484214750] HOST ALERT: INASOL-CyM_ABOGADOS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 5.05 ms
[1484214791] SERVICE ALERT: INASOL-CyM_ABOGADOS;INASOL - CyM_ABOGADOS;CRITICAL;SOFT;1;CRITICAL - IN bandwidth (1240117248325.00%) too high
[1484214811] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 403.34Mb (20.17%), Average OUT: 88.69Mb (4.43%)Total RX: 151133.21Gbits, Total TX: 118100.00Gbits
[1484214852] SERVICE ALERT: INASOL-CyM_ABOGADOS;INASOL - CyM_ABOGADOS;OK;SOFT;2;OK - Average IN: 1.01Mb (1.01%), Average OUT: 85.39Kb (0.09%)Total RX: 69.25Mbits, Total TX: 8.17Mbits
[1484215173] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.89%) too high
[1484215228] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 439.76Mb (21.99%), Average OUT: 74.83Mb (3.74%)Total RX: 151305.52Gbits, Total TX: 118134.36Gbits
[1484215592] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.92%) too high
[1484215653] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 431.59Mb (21.58%), Average OUT: 86.67Mb (4.33%)Total RX: 151495.40Gbits, Total TX: 118169.13Gbits
[1484215698] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.30%) too high
[1484215757] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (87.95%) too high
[1484215818] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (82.95%) too high
[1484215877] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;4;OK - Average IN: 73.82Mb (73.82%), Average OUT: 14.45Mb (14.45%)Total RX: 20478.09Gbits, Total TX: 5108.77Gbits
[1484215998] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.99%) too high
[1484216058] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 69.30Mb (69.30%), Average OUT: 11.42Mb (11.42%)Total RX: 20491.16Gbits, Total TX: 5110.98Gbits
[1484216073] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.85%) too high
[1484216130] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 431.22Mb (21.56%), Average OUT: 80.25Mb (4.01%)Total RX: 151706.84Gbits, Total TX: 118205.44Gbits
[1484216792] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484216837] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (88.87%) too high
[1484216852] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 461.44b (0.00%)Total RX: 8.99Kbits, Total TX: 407.95Mbits
[1484216898] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (96.91%) too high
[1484216957] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (97.07%) too high
[1484216973] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.97%) too high
[1484217018] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (87.79%) too high
[1484217032] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 307.41Mb (15.37%), Average OUT: 85.92Mb (4.30%)Total RX: 152071.45Gbits, Total TX: 118275.45Gbits
[1484217197] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 72.46Mb (72.46%), Average OUT: 12.57Mb (12.57%)Total RX: 20576.06Gbits, Total TX: 5126.79Gbits
[1484217258] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.13%) too high
[1484217318] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 69.16Mb (69.16%), Average OUT: 13.97Mb (13.97%)Total RX: 20584.79Gbits, Total TX: 5128.36Gbits
[1484217389] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (131762457654.80%) too high
[1484217438] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.78%) too high
[1484217451] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 312.51Mb (15.63%), Average OUT: 635.97Mb (31.80%)Total RX: 152214.55Gbits, Total TX: 118344.13Gbits
[1484217498] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (90.25%) too high
[1484217557] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (93.72%) too high
[1484217618] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (88.20%) too high
[1484217812] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.05%) too high
[1484217873] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 361.14Mb (18.06%), Average OUT: 68.67Mb (3.43%)Total RX: 152374.53Gbits, Total TX: 118339.25Gbits
[1484218098] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 70.82Mb (70.82%), Average OUT: 9.24Mb (9.24%)Total RX: 20648.68Gbits, Total TX: 5137.80Gbits
[1484218098] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1484218277] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.15%) too high
[1484218292] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.08%) too high
[1484218337] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 74.65Mb (74.65%), Average OUT: 12.86Mb (12.86%)Total RX: 20666.42Gbits, Total TX: 5140.81Gbits
[1484218352] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 408.39Mb (20.42%), Average OUT: 75.64Mb (3.78%)Total RX: 152562.88Gbits, Total TX: 118375.45Gbits
[1484218397] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (86.20%) too high
[1484218458] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (80.49%) too high
[1484218518] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (76.76%) too high
[1484218577] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (76.45%) too high
[1484218638] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 74.08Mb (74.08%), Average OUT: 9.23Mb (9.23%)Total RX: 20690.13Gbits, Total TX: 5144.02Gbits
[1484218649] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (131762457657.40%) too high
[1484218712] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 386.18Mb (19.31%), Average OUT: 77.03Mb (3.85%)Total RX: 152704.11Gbits, Total TX: 118437.01Gbits
[1484218758] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (79.02%) too high
[1484218771] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671661.36%) too high
[1484218818] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (81.83%) too high
[1484218828] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 372.10Mb (18.60%), Average OUT: 679.00Mb (33.95%)Total RX: 152746.60Gbits, Total TX: 118446.44Gbits
[1484218878] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (78.24%) too high
[1484218938] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (79.08%) too high
[1484218997] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 64.24Mb (64.24%), Average OUT: 9.00Mb (9.00%)Total RX: 20717.38Gbits, Total TX: 5147.53Gbits
[1484218997] SERVICE ALERT: INASOL-BGP01;INASOL [INTERNET REDCA - G0/0/0];WARNING;SOFT;1;WARNING - OUT bandwidth (84.42%) too high
[1484219058] SERVICE ALERT: INASOL-BGP01;INASOL [INTERNET REDCA - G0/0/0];OK;SOFT;2;OK - Average IN: 7.32Mb (7.32%), Average OUT: 22.04Mb (22.04%)Total RX: 22344.89Gbits, Total TX: 29571.73Gbits
[1484219068] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.06, 2.21, 2.96
[1484219118] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.69%) too high
[1484219178] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 74.25Mb (74.25%), Average OUT: 6.88Mb (6.88%)Total RX: 20730.91Gbits, Total TX: 5149.03Gbits
[1484219188] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (134158138690.76%) too high
[1484219238] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (80.23%) too high
[1484219248] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 2.51, 2.80, 3.08
[1484219253] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 886.46Mb (44.32%), Average OUT: 68.40Mb (3.42%)Total RX: 152933.40Gbits, Total TX: 118446.31Gbits
[1484219298] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (76.31%) too high
[1484219309] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (131762457656.69%) too high
[1484219357] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 72.89Mb (72.89%), Average OUT: 7.19Mb (7.19%)Total RX: 20744.60Gbits, Total TX: 5150.41Gbits
[1484219368] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.52, 2.53, 2.97
[1484219372] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 343.74Mb (17.19%), Average OUT: 78.81Mb (3.94%)Total RX: 152940.87Gbits, Total TX: 118455.86Gbits
[1484219477] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (82.06%) too high
[1484219538] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 74.06Mb (74.06%), Average OUT: 11.11Mb (11.11%)Total RX: 20758.34Gbits, Total TX: 5152.07Gbits
[1484219548] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 2.29, 2.82, 3.02
[1484219548] SERVICE FLAPPING ALERT: OS-MONITOREO;Carga OS-MONITOREO;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1484219599] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (76.90%) too high
[1484219609] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.83, 2.73, 2.99
[1484219611] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (127218924621.01%) too high
[1484219657] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (84.26%) too high
[1484219668] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 2.17, 2.87, 3.03
[1484219669] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 441.66Mb (22.08%), Average OUT: 684.66Mb (34.23%)Total RX: 153055.99Gbits, Total TX: 118514.99Gbits
[1484219716] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (82.28%) too high
[1484219776] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;4;OK - Average IN: 74.02Mb (74.02%), Average OUT: 11.57Mb (11.57%)Total RX: 20777.22Gbits, Total TX: 5154.62Gbits
[1484219837] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (77.07%) too high
[1484219852] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484219853] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256212.42%) too high
[1484219896] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (87.16%) too high
[1484219911] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 385.86Mb (19.29%), Average OUT: 84.49Mb (4.22%)Total RX: 153152.03Gbits, Total TX: 118535.39Gbits
[1484219912] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 474.24b (0.00%)Total RX: 8.99Kbits, Total TX: 409.54Mbits
[1484219956] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (96.43%) too high
[1484220016] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (76.37%) too high
[1484220032] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293814.74%) too high
[1484220092] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 341.26Mb (17.06%), Average OUT: 70.56Mb (3.53%)Total RX: 153217.05Gbits, Total TX: 118514.93Gbits
[1484220197] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 71.86Mb (71.86%), Average OUT: 11.45Mb (11.45%)Total RX: 20811.12Gbits, Total TX: 5159.49Gbits
[1484220376] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (82.67%) too high
[1484220392] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293801.18%) too high
[1484220436] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (75.53%) too high
[1484220453] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 332.73Mb (16.64%), Average OUT: 90.68Mb (4.53%)Total RX: 153348.72Gbits, Total TX: 118549.48Gbits
[1484220496] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 73.11Mb (73.11%), Average OUT: 19.59Mb (19.59%)Total RX: 20833.61Gbits, Total TX: 5164.76Gbits
[1484220556] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (81.24%) too high
[1484220617] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (77.93%) too high
[1484220628] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.48, 2.47, 2.97
[1484220676] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (79.60%) too high
[1484220688] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 2.14, 2.71, 3.03
[1484220692] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293815.30%) too high
[1484220736] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (76.38%) too high
[1484220748] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.83, 2.60, 2.98
[1484220752] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (122978293800.54%) too high
[1484220812] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 334.05Mb (16.70%), Average OUT: 92.13Mb (4.61%)Total RX: 153481.16Gbits, Total TX: 118585.31Gbits
[1484220988] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 2.61, 2.87, 3.02
[1484221035] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (91.52%) too high
[1484221155] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (85.92%) too high
[1484221172] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (119011252064.30%) too high
[1484221233] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 380.97Mb (19.05%), Average OUT: 81.69Mb (4.08%)Total RX: 153634.70Gbits, Total TX: 118619.19Gbits
[1484221383] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1484221408] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.33, 2.71, 2.99
[1484221435] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1484221468] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 2.25, 2.94, 3.07
[1484221495] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1484221555] HOST ALERT: INASOL-STEVEN;DOWN;HARD;4;PING CRITICAL - Packet loss = 100%
[1484221635] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 74.69Mb (74.69%), Average OUT: 20.71Mb (20.71%)Total RX: 20927.14Gbits, Total TX: 5188.54Gbits
[1484221888] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 0.93, 2.34, 2.88
[1484221890] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (129450835580.08%) too high
[1484221949] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 380.32Mb (19.02%), Average OUT: 76.46Mb (3.82%)Total RX: 153912.31Gbits, Total TX: 118684.07Gbits
[1484221995] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (76.93%) too high
[1484222055] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (84.62%) too high
[1484222115] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 69.44Mb (69.44%), Average OUT: 17.86Mb (17.86%)Total RX: 20962.51Gbits, Total TX: 5197.88Gbits
[1484222235] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (76.13%) too high
[1484222295] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (76.59%) too high
[1484222356] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (85.74%) too high
[1484222372] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (119011252064.35%) too high
[1484222415] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (76.19%) too high
[1484222433] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 311.62Mb (15.58%), Average OUT: 72.09Mb (3.60%)Total RX: 154091.41Gbits, Total TX: 118720.88Gbits
[1484222475] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 74.64Mb (74.64%), Average OUT: 15.62Mb (15.62%)Total RX: 20990.05Gbits, Total TX: 5204.78Gbits
[1484222595] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (83.69%) too high
[1484222655] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (87.87%) too high
[1484222668] SERVICE FLAPPING ALERT: OS-MONITOREO;Carga OS-MONITOREO;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1484222715] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (83.82%) too high
[1484222775] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (78.54%) too high
[1484222792] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.59%) too high
[1484222852] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 309.19Mb (15.46%), Average OUT: 77.94Mb (3.90%)Total RX: 154208.22Gbits, Total TX: 118753.47Gbits
[1484222912] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484222968] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;OK;HARD;4;OK - load average: 1.53, 1.98, 2.44
[1484222972] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 474.24b (0.00%)Total RX: 9.57Kbits, Total TX: 411.13Mbits
[1484223015] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 74.83Mb (74.83%), Average OUT: 18.36Mb (18.36%)Total RX: 21032.73Gbits, Total TX: 5215.09Gbits
[1484223075] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.72%) too high
[1484223089] HOST ALERT: INASOL-STEVEN;UP;HARD;4;PING OK - Packet loss = 0%, RTA = 11.44 ms
[1484223135] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (83.09%) too high
[1484223195] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (84.79%) too high
[1484223213] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.85%) too high
[1484223256] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (77.02%) too high
[1484223270] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 283.83Mb (14.19%), Average OUT: 72.49Mb (3.62%)Total RX: 154339.53Gbits, Total TX: 118786.24Gbits
[1484223375] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 73.29Mb (73.29%), Average OUT: 25.87Mb (25.87%)Total RX: 21060.95Gbits, Total TX: 5223.93Gbits
[1484223435] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (81.62%) too high
[1484223452] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293812.56%) too high
[1484223495] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 74.75Mb (74.75%), Average OUT: 21.26Mb (21.26%)Total RX: 21070.33Gbits, Total TX: 5226.51Gbits
[1484223513] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 312.98Mb (15.65%), Average OUT: 67.34Mb (3.37%)Total RX: 154415.19Gbits, Total TX: 118804.18Gbits
[1484223689] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (131762457643.06%) too high
[1484223751] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 376.68Mb (18.83%), Average OUT: 83.81Mb (4.19%)Total RX: 154494.39Gbits, Total TX: 118823.39Gbits
[1484223873] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256206.84%) too high
[1484223928] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 926.74Mb (46.34%), Average OUT: 82.08Mb (4.10%)Total RX: 154583.62Gbits, Total TX: 118872.16Gbits
[1484223993] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (113518425057.15%) too high
[1484224049] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 291.19Mb (14.56%), Average OUT: 79.08Mb (3.95%)Total RX: 154584.54Gbits, Total TX: 118881.86Gbits
[1484224112] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (117122184572.10%) too high
[1484224173] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 296.54Mb (14.83%), Average OUT: 97.60Mb (4.88%)Total RX: 154621.86Gbits, Total TX: 118858.98Gbits
[1484224471] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (119011252072.30%) too high
[1484224515] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1484224533] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 856.41Mb (42.82%), Average OUT: 70.15Mb (3.51%)Total RX: 154754.16Gbits, Total TX: 118891.24Gbits
[1484224575] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (80.16%) too high
[1484224589] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (131762457653.47%) too high
[1484224636] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (82.36%) too high
[1484224652] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 857.90Mb (42.90%), Average OUT: 76.95Mb (3.85%)Total RX: 154790.42Gbits, Total TX: 118900.58Gbits
[1484224695] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (91.14%) too high
[1484224712] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293810.85%) too high
[1484224754] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (83.50%) too high
[1484224773] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 282.66Mb (14.13%), Average OUT: 68.31Mb (3.42%)Total RX: 154791.01Gbits, Total TX: 118908.83Gbits
[1484224994] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 72.96Mb (72.96%), Average OUT: 24.91Mb (24.91%)Total RX: 21177.93Gbits, Total TX: 5251.85Gbits
[1484225055] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (91.47%) too high
[1484225069] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (134158138702.11%) too high
[1484225069] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;1;WARNING - load average: 3.82, 3.14, 2.60
[1484225114] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (84.80%) too high
[1484225130] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;2;WARNING - load average: 3.11, 3.14, 2.64
[1484225134] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 287.61Mb (14.38%), Average OUT: 89.80Mb (4.49%)Total RX: 154894.88Gbits, Total TX: 118972.93Gbits
[1484225174] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 72.56Mb (72.56%), Average OUT: 22.29Mb (22.29%)Total RX: 21192.87Gbits, Total TX: 5256.07Gbits
[1484225189] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;3;WARNING - load average: 3.79, 3.55, 2.83
[1484225249] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 3.45, 3.66, 2.93
[1484225294] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (83.63%) too high
[1484225354] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (80.57%) too high
[1484225373] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.20%) too high
[1484225414] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 72.65Mb (72.65%), Average OUT: 14.65Mb (14.65%)Total RX: 21211.46Gbits, Total TX: 5260.32Gbits
[1484225434] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 315.47Mb (15.77%), Average OUT: 92.56Mb (4.63%)Total RX: 154986.92Gbits, Total TX: 118963.19Gbits
[1484225493] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1484225548] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 3.95, 3.77, 3.15
[1484225574] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1484225594] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (85.72%) too high
[1484225635] HOST ALERT: INASOL-STEVEN;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1484225654] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (84.37%) too high
[1484225695] HOST ALERT: INASOL-STEVEN;DOWN;HARD;4;PING CRITICAL - Packet loss = 100%
[1484225714] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (86.38%) too high
[1484225774] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (88.42%) too high
[1484225792] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.75%) too high
[1484225848] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 2.98, 3.06, 2.99
[1484225852] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 341.06Mb (17.05%), Average OUT: 80.75Mb (4.04%)Total RX: 155129.96Gbits, Total TX: 118995.47Gbits
[1484226014] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (94.69%) too high
[1484226014] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1484226134] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (87.39%) too high
[1484226211] HOST ALERT: INASOL-STEVEN;UP;HARD;4;PING OK - Packet loss = 37%, RTA = 10.45 ms
[1484226212] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293810.22%) too high
[1484226271] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (125062671660.58%) too high
[1484226317] SERVICE ALERT: INASOL-BGP01;INASOL [INTERNET REDCA - G0/0/0];WARNING;SOFT;1;WARNING - OUT bandwidth (88.16%) too high
[1484226332] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (120962256207.41%) too high
[1484226378] SERVICE ALERT: INASOL-BGP01;INASOL [INTERNET REDCA - G0/0/0];OK;SOFT;2;OK - Average IN: 14.17Mb (14.17%), Average OUT: 68.37Mb (68.37%)Total RX: 22417.11Gbits, Total TX: 29691.67Gbits
[1484226392] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;4;OK - Average IN: 299.90Mb (14.99%), Average OUT: 68.96Mb (3.45%)Total RX: 155306.26Gbits, Total TX: 119034.44Gbits
[1484226493] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.07%) too high
[1484226553] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (79.32%) too high
[1484226692] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.53%) too high
[1484226752] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (122978293813.40%) too high
[1484226794] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 70.86Mb (70.86%), Average OUT: 13.33Mb (13.33%)Total RX: 21326.45Gbits, Total TX: 5281.22Gbits
[1484226812] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484226813] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 375.42Mb (18.77%), Average OUT: 76.55Mb (3.83%)Total RX: 155451.57Gbits, Total TX: 119065.99Gbits
[1484226872] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 474.56b (0.00%)Total RX: 9.57Kbits, Total TX: 413.15Mbits
[1484226972] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.91%) too high
[1484227032] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 70.13Mb (70.13%), Average OUT: 13.05Mb (13.05%)Total RX: 21343.34Gbits, Total TX: 5284.34Gbits
[1484227288] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 4.04, 3.53, 3.09
[1484227348] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 2.53, 3.19, 3.00
[1484227392] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (79.68%) too high
[1484227452] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (85.56%) too high
[1484227472] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293817.44%) too high
[1484227512] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (81.96%) too high
[1484227532] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 411.67Mb (20.58%), Average OUT: 78.17Mb (3.91%)Total RX: 155732.86Gbits, Total TX: 119158.16Gbits
[1484227572] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (85.24%) too high
[1484227588] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 4.23, 3.35, 3.08
[1484227593] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256197.34%) too high
[1484227633] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.50%) too high
[1484227652] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 1.05Gb (52.41%), Average OUT: 81.48Mb (4.07%)Total RX: 155820.13Gbits, Total TX: 119134.05Gbits
[1484227692] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (85.93%) too high
[1484227712] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293815.60%) too high
[1484227772] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 377.39Mb (18.87%), Average OUT: 76.69Mb (3.83%)Total RX: 155831.82Gbits, Total TX: 119143.61Gbits
[1484227811] SERVICE ALERT: FTTH-SABANA01;INASOL [INTERNET REDCA - G1/8];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 8 status: 2)
[1484227812] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 72.47Mb (72.47%), Average OUT: 10.97Mb (10.97%)Total RX: 21402.45Gbits, Total TX: 5296.11Gbits
[1484227871] SERVICE ALERT: FTTH-SABANA01;INASOL [INTERNET REDCA - G1/8];OK;SOFT;2;OK - Average IN: 7.78Mb (0.78%), Average OUT: 7.74Mb (0.77%)Total RX: 4368.32Gbits, Total TX: 47434.00Gbits
[1484227872] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.47%) too high
[1484227932] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (80.06%) too high
[1484227993] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (84.30%) too high
[1484228012] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.04%) too high
[1484228052] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (80.56%) too high
[1484228068] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.33, 2.78, 3.00
[1484228068] SERVICE FLAPPING ALERT: OS-MONITOREO;Carga OS-MONITOREO;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1484228072] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 423.27Mb (21.16%), Average OUT: 84.35Mb (4.22%)Total RX: 155951.76Gbits, Total TX: 119164.81Gbits
[1484228128] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 2.04, 2.85, 3.02
[1484228188] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.36, 2.60, 2.93
[1484228292] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 70.68Mb (70.68%), Average OUT: 9.22Mb (9.22%)Total RX: 21439.96Gbits, Total TX: 5301.26Gbits
[1484228472] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (77.28%) too high
[1484228493] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.75%) too high
[1484228532] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (75.98%) too high
[1484228552] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (125062671677.29%) too high
[1484228591] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 71.68Mb (71.68%), Average OUT: 16.43Mb (16.43%)Total RX: 21462.01Gbits, Total TX: 5305.12Gbits
[1484228613] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 400.56Mb (20.03%), Average OUT: 76.86Mb (3.84%)Total RX: 156177.10Gbits, Total TX: 119206.95Gbits
[1484228651] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (78.60%) too high
[1484228711] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (80.08%) too high
[1484228771] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 71.46Mb (71.46%), Average OUT: 17.68Mb (17.68%)Total RX: 21475.81Gbits, Total TX: 5307.90Gbits
[1484228891] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (82.53%) too high
[1484228908] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 3.31, 3.37, 3.09
[1484228913] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.74%) too high
[1484228951] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (81.23%) too high
[1484228968] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 2.14, 3.04, 2.99
[1484228973] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 358.32Mb (17.92%), Average OUT: 80.34Mb (4.02%)Total RX: 156323.85Gbits, Total TX: 119235.04Gbits
[1484229011] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (78.11%) too high
[1484229071] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;4;OK - Average IN: 69.26Mb (69.26%), Average OUT: 16.07Mb (16.07%)Total RX: 21498.57Gbits, Total TX: 5312.47Gbits
[1484229088] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 3.00, 3.15, 3.04
[1484229208] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 2.20, 2.96, 3.00
[1484229251] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.28%) too high
[1484229268] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 2.96, 3.20, 3.09
[1484229311] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (85.96%) too high
[1484229332] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293814.04%) too high
[1484229371] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;3;WARNING - IN bandwidth (85.64%) too high
[1484229392] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 388.76Mb (19.44%), Average OUT: 94.77Mb (4.74%)Total RX: 156483.13Gbits, Total TX: 119271.42Gbits
[1484229431] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (90.58%) too high
[1484229491] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (89.36%) too high
[1484229752] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.50%) too high
[1484229791] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (92.07%) too high
[1484229812] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 371.75Mb (18.59%), Average OUT: 74.77Mb (3.74%)Total RX: 156635.27Gbits, Total TX: 119302.19Gbits
[1484229851] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (83.22%) too high
[1484230031] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 20.12Mb (20.12%), Average OUT: 5.28Mb (5.28%)Total RX: 21575.36Gbits, Total TX: 5323.33Gbits
[1484230212] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (80.31%) too high
[1484230271] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (87.64%) too high
[1484230332] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;SOFT;3;CRITICAL - IN bandwidth (94.09%) too high
[1484230391] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.98%) too high
[1484230408] SERVICE FLAPPING ALERT: OS-MONITOREO;Carga OS-MONITOREO;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1484230412] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.86%) too high
[1484230451] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (88.25%) too high
[1484230472] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 55.75Mb (2.79%), Average OUT: 624.32Mb (31.22%)Total RX: 156710.91Gbits, Total TX: 119371.10Gbits
[1484230512] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (93.36%) too high
[1484230713] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484230751] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (81.87%) too high
[1484230772] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 464.88b (0.00%)Total RX: 9.57Kbits, Total TX: 415.18Mbits
[1484230949] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.54%) too high
[1484231009] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 223.46Mb (11.17%), Average OUT: 64.94Mb (3.25%)Total RX: 156792.51Gbits, Total TX: 119370.01Gbits
[1484231051] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (90.75%) too high
[1484231369] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (127218924624.57%) too high
[1484231412] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (89.17%) too high
[1484231433] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 181.19Mb (9.06%), Average OUT: 81.62Mb (4.08%)Total RX: 156872.55Gbits, Total TX: 119433.37Gbits
[1484231488] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (134158138689.99%) too high
[1484231553] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484231553] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 180.65Mb (9.03%), Average OUT: 57.05Mb (2.85%)Total RX: 156894.56Gbits, Total TX: 119406.39Gbits
[1484231612] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 461.44b (0.00%)Total RX: 9.57Kbits, Total TX: 415.61Mbits
[1484231732] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293807.05%) too high
[1484231792] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 240.29Mb (12.01%), Average OUT: 61.28Mb (3.06%)Total RX: 156948.30Gbits, Total TX: 119422.07Gbits
[1484231973] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.64%) too high
[1484232029] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 229.91Mb (11.50%), Average OUT: 69.90Mb (3.50%)Total RX: 157004.42Gbits, Total TX: 119438.60Gbits
[1484232132] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];CRITICAL;HARD;4;CRITICAL - IN bandwidth (92.52%) too high
[1484232191] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;HARD;4;WARNING - IN bandwidth (89.90%) too high
[1484232312] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;HARD;4;OK - Average IN: 68.41Mb (68.41%), Average OUT: 13.87Mb (13.87%)Total RX: 21768.35Gbits, Total TX: 5354.19Gbits
[1484232513] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256197.45%) too high
[1484232569] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 194.51Mb (9.73%), Average OUT: 70.97Mb (3.55%)Total RX: 157110.77Gbits, Total TX: 119477.61Gbits
[1484232672] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (83.43%) too high
[1484232731] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (75.31%) too high
[1484232791] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 65.84Mb (65.84%), Average OUT: 15.95Mb (15.95%)Total RX: 21802.60Gbits, Total TX: 5360.95Gbits
[1484232933] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (113518425045.53%) too high
[1484232989] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 67.06Mb (3.35%), Average OUT: 52.85Mb (2.64%)Total RX: 157178.77Gbits, Total TX: 119506.72Gbits
[1484233288] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (134158138691.29%) too high
[1484233352] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 75.94Mb (3.80%), Average OUT: 75.69Mb (3.78%)Total RX: 157204.33Gbits, Total TX: 119563.70Gbits
[1484233451] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.29%) too high
[1484233511] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;2;WARNING - IN bandwidth (80.10%) too high
[1484233532] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.88%) too high
[1484233571] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;3;OK - Average IN: 64.16Mb (64.16%), Average OUT: 25.35Mb (25.35%)Total RX: 21855.58Gbits, Total TX: 5375.67Gbits
[1484233592] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 83.48Mb (4.17%), Average OUT: 53.91Mb (2.70%)Total RX: 157222.49Gbits, Total TX: 119542.03Gbits
[1484233631] SERVICE FLAPPING ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1484233691] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (83.60%) too high
[1484233751] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 71.23Mb (71.23%), Average OUT: 10.06Mb (10.06%)Total RX: 21869.34Gbits, Total TX: 5378.04Gbits
[1484234052] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];WARNING;SOFT;1;WARNING - IN bandwidth (75.82%) too high
[1484234112] SERVICE ALERT: FTTH-AMERICAN_DATA01;AMERICAN DATA [PLAZA MAYOR - F0/8];OK;SOFT;2;OK - Average IN: 68.83Mb (68.83%), Average OUT: 10.63Mb (10.63%)Total RX: 21895.18Gbits, Total TX: 5382.17Gbits
[1484234194] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (119011252064.45%) too high
[1484234248] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (136642548664.84%) too high
[1484234313] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 66.52Mb (3.33%), Average OUT: 54.70Mb (2.73%)Total RX: 157276.49Gbits, Total TX: 119577.60Gbits
[1484234492] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293803.68%) too high
[1484234552] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 172.32Mb (8.62%), Average OUT: 58.82Mb (2.94%)Total RX: 157316.16Gbits, Total TX: 119592.43Gbits
[1484234790] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (129450835577.94%) too high
[1484234852] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 184.95Mb (9.25%), Average OUT: 63.87Mb (3.19%)Total RX: 157375.17Gbits, Total TX: 119611.22Gbits
[1484235329] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (131762457641.70%) too high
[1484235389] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 202.50Mb (10.12%), Average OUT: 55.07Mb (2.75%)Total RX: 157472.82Gbits, Total TX: 119645.00Gbits
[1484235452] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484235512] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 465.76b (0.00%)Total RX: 9.57Kbits, Total TX: 417.64Mbits
[1484235873] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (115292150436.89%) too high
[1484235928] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 230.00Mb (11.50%), Average OUT: 67.05Mb (3.35%)Total RX: 157581.50Gbits, Total TX: 119678.76Gbits
[1484236107] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671669.59%) too high
[1484236168] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 245.11Mb (12.26%), Average OUT: 56.00Mb (2.80%)Total RX: 157637.37Gbits, Total TX: 119692.93Gbits
[1484236528] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293808.46%) too high
[1484236588] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 248.86Mb (12.44%), Average OUT: 53.88Mb (2.69%)Total RX: 157740.90Gbits, Total TX: 119749.99Gbits
[1484237008] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.51%) too high
[1484237068] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 182.47Mb (9.12%), Average OUT: 88.31Mb (4.42%)Total RX: 157837.32Gbits, Total TX: 119750.86Gbits
[1484237486] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.12%) too high
[1484237546] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 183.29Mb (9.16%), Average OUT: 62.62Mb (3.13%)Total RX: 157937.10Gbits, Total TX: 119782.23Gbits
[1484238025] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.79%) too high
[1484238086] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 216.26Mb (10.81%), Average OUT: 58.16Mb (2.91%)Total RX: 158043.01Gbits, Total TX: 119817.56Gbits
[1484238625] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.17%) too high
[1484238684] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 200.86Mb (10.04%), Average OUT: 77.65Mb (3.88%)Total RX: 158170.39Gbits, Total TX: 119854.87Gbits
[1484239044] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256201.83%) too high
[1484239104] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 184.77Mb (9.24%), Average OUT: 60.32Mb (3.02%)Total RX: 158246.96Gbits, Total TX: 119878.54Gbits
[1484239352] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484239412] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 470.8b (0.00%)Total RX: 9.57Kbits, Total TX: 419.66Mbits
[1484239463] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671664.92%) too high
[1484239524] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 188.18Mb (9.41%), Average OUT: 64.45Mb (3.22%)Total RX: 158314.68Gbits, Total TX: 119937.41Gbits
[1484239764] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.90%) too high
[1484239823] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 165.33Mb (8.27%), Average OUT: 69.84Mb (3.49%)Total RX: 158371.40Gbits, Total TX: 119922.25Gbits
[1484240243] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671663.82%) too high
[1484240304] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 158.60Mb (7.93%), Average OUT: 78.33Mb (3.92%)Total RX: 158448.83Gbits, Total TX: 119957.24Gbits
[1484240723] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.92%) too high
[1484240784] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 151.41Mb (7.57%), Average OUT: 627.50Mb (31.38%)Total RX: 158520.63Gbits, Total TX: 120027.01Gbits
[1484241143] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.53%) too high
[1484241204] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 170.83Mb (8.54%), Average OUT: 64.35Mb (3.22%)Total RX: 158591.67Gbits, Total TX: 120023.08Gbits
[1484241624] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.81%) too high
[1484241684] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 127.28Mb (6.36%), Average OUT: 79.06Mb (3.95%)Total RX: 158660.12Gbits, Total TX: 120058.72Gbits
[1484242044] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293800.48%) too high
[1484242104] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 137.65Mb (6.88%), Average OUT: 83.75Mb (4.19%)Total RX: 158718.73Gbits, Total TX: 120092.77Gbits
[1484242412] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484242472] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 475.36b (0.00%)Total RX: 9.57Kbits, Total TX: 421.26Mbits
[1484242523] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.90%) too high
[1484242584] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 130.65Mb (6.53%), Average OUT: 65.97Mb (3.30%)Total RX: 158781.79Gbits, Total TX: 120128.99Gbits
[1484243004] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.37%) too high
[1484243064] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 118.84Mb (5.94%), Average OUT: 74.27Mb (3.71%)Total RX: 158835.04Gbits, Total TX: 120163.27Gbits
[1484243544] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.64%) too high
[1484243604] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 113.30Mb (5.66%), Average OUT: 64.03Mb (3.20%)Total RX: 158894.04Gbits, Total TX: 120196.71Gbits
[1484243843] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671662.58%) too high
[1484243904] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 104.82Mb (5.24%), Average OUT: 46.99Mb (2.35%)Total RX: 158925.48Gbits, Total TX: 120247.33Gbits
[1484244204] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.37%) too high
[1484244263] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 99.38Mb (4.97%), Average OUT: 50.84Mb (2.54%)Total RX: 158961.67Gbits, Total TX: 120230.52Gbits
[1484244803] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.06%) too high
[1484244864] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 102.31Mb (5.12%), Average OUT: 59.19Mb (2.96%)Total RX: 159019.18Gbits, Total TX: 120265.09Gbits
[1484245884] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.21%) too high
[1484245944] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 117.42Mb (5.87%), Average OUT: 56.71Mb (2.84%)Total RX: 159127.89Gbits, Total TX: 120333.20Gbits
[1484245948] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 0.46, 1.40, 2.84
[1484246244] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256198.03%) too high
[1484246304] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 96.59Mb (4.83%), Average OUT: 60.43Mb (3.02%)Total RX: 159165.09Gbits, Total TX: 120385.41Gbits
[1484246483] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.54%) too high
[1484246544] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 664.16Mb (33.21%), Average OUT: 62.45Mb (3.12%)Total RX: 159221.55Gbits, Total TX: 120366.80Gbits
[1484246604] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293800.75%) too high
[1484246663] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 102.73Mb (5.14%), Average OUT: 63.63Mb (3.18%)Total RX: 159198.83Gbits, Total TX: 120374.12Gbits
[1484246848] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 3.00, 3.23, 3.08
[1484247024] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.42%) too high
[1484247083] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 97.05Mb (4.85%), Average OUT: 61.47Mb (3.07%)Total RX: 159239.75Gbits, Total TX: 120401.45Gbits
[1484247208] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.64, 2.75, 2.97
[1484247324] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256196.63%) too high
[1484247384] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 78.99Mb (3.95%), Average OUT: 54.99Mb (2.75%)Total RX: 159264.83Gbits, Total TX: 120419.30Gbits
[1484247388] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 3.94, 3.07, 3.04
[1484247688] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.80, 2.77, 2.97
[1484247688] SERVICE FLAPPING ALERT: OS-MONITOREO;Carga OS-MONITOREO;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1484247804] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293799.26%) too high
[1484247863] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 55.97Mb (2.80%), Average OUT: 74.54Mb (3.73%)Total RX: 159298.23Gbits, Total TX: 120484.95Gbits
[1484248103] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.55%) too high
[1484248164] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 74.43Mb (3.72%), Average OUT: 63.54Mb (3.18%)Total RX: 159320.78Gbits, Total TX: 120470.67Gbits
[1484248284] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293799.34%) too high
[1484248343] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 56.88Mb (2.84%), Average OUT: 648.31Mb (32.42%)Total RX: 159332.04Gbits, Total TX: 120516.72Gbits
[1484248588] SERVICE FLAPPING ALERT: OS-MONITOREO;Carga OS-MONITOREO;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1484248644] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.48%) too high
[1484248703] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 662.99Mb (33.15%), Average OUT: 55.51Mb (2.78%)Total RX: 159394.04Gbits, Total TX: 120503.86Gbits
[1484248764] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256196.24%) too high
[1484248824] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 66.98Mb (3.35%), Average OUT: 65.79Mb (3.29%)Total RX: 159367.80Gbits, Total TX: 120511.41Gbits
[1484249008] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;OK;HARD;4;OK - load average: 1.15, 1.94, 2.43
[1484249244] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.12%) too high
[1484249304] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 76.43Mb (3.82%), Average OUT: 64.47Mb (3.22%)Total RX: 159402.48Gbits, Total TX: 120541.41Gbits
[1484249664] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293801.25%) too high
[1484249724] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 91.43Mb (4.57%), Average OUT: 634.65Mb (31.73%)Total RX: 159440.64Gbits, Total TX: 120598.52Gbits
[1484249843] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.33%) too high
[1484249904] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 81.88Mb (4.09%), Average OUT: 623.24Mb (31.16%)Total RX: 159455.82Gbits, Total TX: 120608.05Gbits
[1484250324] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.76%) too high
[1484250383] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 645.41Mb (32.27%), Average OUT: 70.60Mb (3.53%)Total RX: 159530.45Gbits, Total TX: 120606.89Gbits
[1484250444] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256195.91%) too high
[1484250503] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 71.48Mb (3.57%), Average OUT: 78.34Mb (3.92%)Total RX: 159503.99Gbits, Total TX: 120616.23Gbits
[1484250864] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.90%) too high
[1484250924] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 55.91Mb (2.80%), Average OUT: 48.45Mb (2.42%)Total RX: 159530.13Gbits, Total TX: 120642.82Gbits
[1484251288] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;1;WARNING - load average: 4.14, 2.93, 2.33
[1484251348] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;OK;SOFT;2;OK - load average: 3.45, 2.98, 2.39
[1484251528] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;1;WARNING - load average: 3.81, 3.09, 2.52
[1484251588] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;2;WARNING - load average: 3.08, 3.02, 2.53
[1484251593] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484251643] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671659.48%) too high
[1484251648] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;SOFT;3;CRITICAL - load average: 7.20, 3.93, 2.86
[1484251652] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 474.24b (0.00%)Total RX: 9.57Kbits, Total TX: 426.03Mbits
[1484251704] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 53.08Mb (2.65%), Average OUT: 62.65Mb (3.13%)Total RX: 159572.04Gbits, Total TX: 120723.94Gbits
[1484251708] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 5.47, 3.99, 2.95
[1484251769] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 5.66, 4.23, 3.09
[1484252004] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.89%) too high
[1484252063] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 75.97Mb (3.80%), Average OUT: 68.01Mb (3.40%)Total RX: 159591.08Gbits, Total TX: 120713.43Gbits
[1484252424] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.80%) too high
[1484252483] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 47.02Mb (2.35%), Average OUT: 74.90Mb (3.74%)Total RX: 159615.27Gbits, Total TX: 120746.15Gbits
[1484252904] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.52%) too high
[1484252964] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 58.46Mb (2.92%), Average OUT: 65.48Mb (3.27%)Total RX: 159639.32Gbits, Total TX: 120779.83Gbits
[1484253504] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293799.07%) too high
[1484253563] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 60.32Mb (3.02%), Average OUT: 58.95Mb (2.95%)Total RX: 159675.54Gbits, Total TX: 120815.45Gbits
[1484253984] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293800.23%) too high
[1484254044] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - OUT bandwidth (122978293798.80%) too high
[1484254103] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 83.84Mb (4.19%), Average OUT: 58.62Mb (2.93%)Total RX: 159717.11Gbits, Total TX: 120848.52Gbits
[1484254652] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484254704] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.72%) too high
[1484254712] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 471.12b (0.00%)Total RX: 9.57Kbits, Total TX: 427.62Mbits
[1484254763] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 65.05Mb (3.25%), Average OUT: 55.84Mb (2.79%)Total RX: 159765.25Gbits, Total TX: 120883.23Gbits
[1484254944] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256196.02%) too high
[1484255004] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 67.16Mb (3.36%), Average OUT: 629.97Mb (31.50%)Total RX: 159780.51Gbits, Total TX: 120931.74Gbits
[1484255244] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.20%) too high
[1484255304] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 61.94Mb (3.10%), Average OUT: 63.64Mb (3.18%)Total RX: 159801.69Gbits, Total TX: 120916.60Gbits
[1484255783] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.23%) too high
[1484255843] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 60.03Mb (3.00%), Average OUT: 57.21Mb (2.86%)Total RX: 159832.41Gbits, Total TX: 120949.96Gbits
[1484256384] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.46%) too high
[1484256444] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 48.21Mb (2.41%), Average OUT: 54.13Mb (2.71%)Total RX: 159869.98Gbits, Total TX: 120983.90Gbits
[1484256983] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.27%) too high
[1484257044] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 55.03Mb (2.75%), Average OUT: 58.56Mb (2.93%)Total RX: 159902.56Gbits, Total TX: 121019.12Gbits
[1484257223] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671659.76%) too high
[1484257284] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 68.92Mb (3.45%), Average OUT: 55.85Mb (2.79%)Total RX: 159916.61Gbits, Total TX: 121067.80Gbits
[1484257584] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.79%) too high
[1484257643] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 96.09Mb (4.80%), Average OUT: 53.51Mb (2.68%)Total RX: 159946.99Gbits, Total TX: 121054.14Gbits
[1484258183] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.94%) too high
[1484258244] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (120962256195.81%) too high
[1484258304] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 50.07Mb (2.50%), Average OUT: 66.19Mb (3.31%)Total RX: 159985.45Gbits, Total TX: 121095.93Gbits
[1484258664] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.59%) too high
[1484258723] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 61.74Mb (3.09%), Average OUT: 60.80Mb (3.04%)Total RX: 160006.74Gbits, Total TX: 121121.99Gbits
[1484259131] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [CHESSYS - F0/13];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10013 status: 2)
[1484259192] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [CHESSYS - F0/13];CRITICAL;SOFT;2;CRITICAL: Interface is 'down' (interface: 10013 status: 2)
[1484259251] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [CHESSYS - F0/13];OK;SOFT;3;OK - Average IN: 3.90Kb (0.00%), Average OUT: 4.62Kb (0.00%)Total RX: 49.54Gbits, Total TX: 399.57Gbits
[1484259264] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.55%) too high
[1484259324] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 54.96Mb (2.75%), Average OUT: 70.72Mb (3.54%)Total RX: 160039.13Gbits, Total TX: 121157.07Gbits
[1484259932] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484259984] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.23%) too high
[1484259992] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 470b (0.00%)Total RX: 9.57Kbits, Total TX: 430.36Mbits
[1484260043] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 60.31Mb (3.02%), Average OUT: 64.54Mb (3.23%)Total RX: 160079.55Gbits, Total TX: 121193.90Gbits
[1484260463] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.56%) too high
[1484260524] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 57.64Mb (2.88%), Average OUT: 65.53Mb (3.28%)Total RX: 160104.55Gbits, Total TX: 121226.33Gbits
[1484260764] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256195.87%) too high
[1484260823] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 56.30Mb (2.82%), Average OUT: 72.75Mb (3.64%)Total RX: 160121.32Gbits, Total TX: 121278.89Gbits
[1484261004] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.25%) too high
[1484261063] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 54.96Mb (2.75%), Average OUT: 51.08Mb (2.55%)Total RX: 160137.91Gbits, Total TX: 121259.64Gbits
[1484261363] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671659.68%) too high
[1484261424] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 54.56Mb (2.73%), Average OUT: 59.82Mb (2.99%)Total RX: 160156.48Gbits, Total TX: 121314.26Gbits
[1484261604] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.08%) too high
[1484261664] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 58.22Mb (2.91%), Average OUT: 52.84Mb (2.64%)Total RX: 160169.37Gbits, Total TX: 121295.09Gbits
[1484262023] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671659.61%) too high
[1484262084] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 60.34Mb (3.02%), Average OUT: 81.45Mb (4.07%)Total RX: 160191.05Gbits, Total TX: 121356.50Gbits
[1484262144] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.31%) too high
[1484262203] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 50.03Mb (2.50%), Average OUT: 645.03Mb (32.25%)Total RX: 160197.38Gbits, Total TX: 121364.05Gbits
[1484262684] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256195.19%) too high
[1484262744] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 55.18Mb (2.76%), Average OUT: 61.19Mb (3.06%)Total RX: 160224.51Gbits, Total TX: 121366.50Gbits
[1484263223] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.69%) too high
[1484263284] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 42.18Mb (2.11%), Average OUT: 73.63Mb (3.68%)Total RX: 160250.67Gbits, Total TX: 121401.69Gbits
[1484263404] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256195.06%) too high
[1484263463] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 40.92Mb (2.05%), Average OUT: 63.83Mb (3.19%)Total RX: 160258.19Gbits, Total TX: 121413.39Gbits
[1484263704] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.18%) too high
[1484263763] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 40.09Mb (2.00%), Average OUT: 65.92Mb (3.30%)Total RX: 160270.69Gbits, Total TX: 121432.64Gbits
[1484264243] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.95%) too high
[1484264304] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 33.44Mb (1.67%), Average OUT: 45.61Mb (2.28%)Total RX: 160292.71Gbits, Total TX: 121464.87Gbits
[1484264844] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.69%) too high
[1484264903] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 42.35Mb (2.12%), Average OUT: 52.52Mb (2.63%)Total RX: 160317.44Gbits, Total TX: 121501.06Gbits
[1484265212] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484265272] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 476.24b (0.00%)Total RX: 9.57Kbits, Total TX: 433.11Mbits
[1484265443] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.22%) too high
[1484265504] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 36.66Mb (1.83%), Average OUT: 630.05Mb (31.50%)Total RX: 160346.55Gbits, Total TX: 121572.00Gbits
[1484265804] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256195.44%) too high
[1484265863] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 46.50Mb (2.32%), Average OUT: 51.76Mb (2.59%)Total RX: 160361.61Gbits, Total TX: 121593.71Gbits
[1484265983] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.21%) too high
[1484266044] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 39.59Mb (1.98%), Average OUT: 56.75Mb (2.84%)Total RX: 160370.03Gbits, Total TX: 121570.03Gbits
[1484266583] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.09%) too high
[1484266644] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 46.67Mb (2.33%), Average OUT: 67.17Mb (3.36%)Total RX: 160397.48Gbits, Total TX: 121604.78Gbits
[1484267183] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.85%) too high
[1484267244] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (120962256194.57%) too high
[1484267304] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 36.87Mb (1.84%), Average OUT: 636.49Mb (31.82%)Total RX: 160429.34Gbits, Total TX: 121678.11Gbits
[1484267432] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484267492] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 471.12b (0.00%)Total RX: 9.57Kbits, Total TX: 434.26Mbits
[1484267724] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256196.50%) too high
[1484267784] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 40.15Mb (2.01%), Average OUT: 66.22Mb (3.31%)Total RX: 160447.59Gbits, Total TX: 121673.30Gbits
[1484268144] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256194.55%) too high
[1484268203] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 44.07Mb (2.20%), Average OUT: 62.18Mb (3.11%)Total RX: 160463.14Gbits, Total TX: 121734.75Gbits
[1484268264] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.98%) too high
[1484268323] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 35.89Mb (1.79%), Average OUT: 60.52Mb (3.03%)Total RX: 160467.33Gbits, Total TX: 121707.74Gbits
[1484268804] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.84%) too high
[1484268863] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 31.53Mb (1.58%), Average OUT: 659.40Mb (32.97%)Total RX: 160486.23Gbits, Total TX: 121779.11Gbits
[1484269164] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293798.05%) too high
[1484269223] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 39.41Mb (1.97%), Average OUT: 63.55Mb (3.18%)Total RX: 160498.83Gbits, Total TX: 121804.74Gbits
[1484269283] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.35%) too high
[1484269344] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 38.54Mb (1.93%), Average OUT: 67.30Mb (3.36%)Total RX: 160503.50Gbits, Total TX: 121778.39Gbits
[1484269703] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.92%) too high
[1484269764] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 36.87Mb (1.84%), Average OUT: 75.14Mb (3.76%)Total RX: 160520.66Gbits, Total TX: 121811.71Gbits
[1484270003] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.63%) too high
[1484270064] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 36.84Mb (1.84%), Average OUT: 61.67Mb (3.08%)Total RX: 160531.75Gbits, Total TX: 121832.94Gbits
[1484270184] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.64%) too high
[1484270243] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 35.59Mb (1.78%), Average OUT: 51.20Mb (2.56%)Total RX: 160538.52Gbits, Total TX: 121842.99Gbits
[1484270904] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.60%) too high
[1484270963] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;2;CRITICAL - IN bandwidth (125062671658.76%) too high
[1484271024] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;3;OK - Average IN: 30.39Mb (1.52%), Average OUT: 620.78Mb (31.04%)Total RX: 160565.40Gbits, Total TX: 121916.05Gbits
[1484271624] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.76%) too high
[1484271683] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 38.82Mb (1.94%), Average OUT: 57.46Mb (2.87%)Total RX: 160589.25Gbits, Total TX: 121914.74Gbits
[1484271923] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.96%) too high
[1484271984] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 46.83Mb (2.34%), Average OUT: 34.35Mb (1.72%)Total RX: 160601.00Gbits, Total TX: 121963.42Gbits
[1484272168] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 1.00, 1.65, 2.95
[1484272283] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.81%) too high
[1484272343] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 39.97Mb (2.00%), Average OUT: 54.62Mb (2.73%)Total RX: 160615.18Gbits, Total TX: 121947.50Gbits
[1484272712] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484272764] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256194.97%) too high
[1484272773] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 465.76b (0.00%)Total RX: 9.57Kbits, Total TX: 437.00Mbits
[1484272823] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 41.34Mb (2.07%), Average OUT: 54.98Mb (2.75%)Total RX: 160635.05Gbits, Total TX: 122007.56Gbits
[1484272943] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.05%) too high
[1484272949] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;OK;HARD;4;OK - load average: 0.33, 1.64, 2.39
[1484273004] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 27.90Mb (1.40%), Average OUT: 50.51Mb (2.53%)Total RX: 160641.15Gbits, Total TX: 121983.18Gbits
[1484273783] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.67%) too high
[1484273844] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 35.56Mb (1.78%), Average OUT: 55.85Mb (2.79%)Total RX: 160670.14Gbits, Total TX: 122063.31Gbits
[1484274144] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.42%) too high
[1484274203] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 40.53Mb (2.03%), Average OUT: 51.94Mb (2.60%)Total RX: 160683.51Gbits, Total TX: 122049.08Gbits
[1484274804] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.89%) too high
[1484274863] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 38.02Mb (1.90%), Average OUT: 65.01Mb (3.25%)Total RX: 160708.16Gbits, Total TX: 122084.52Gbits
[1484274932] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 10006 status: 2)
[1484274992] SERVICE ALERT: FTTH-SANTA_ANA01;AMERICAN DATA [PRISMAN BUSSINES - F0/6];OK;SOFT;2;OK - Average IN: 0b (0.00%), Average OUT: 470b (0.00%)Total RX: 9.57Kbits, Total TX: 438.16Mbits
[1484275344] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293798.93%) too high
[1484275404] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 610.53Mb (30.53%), Average OUT: 62.82Mb (3.14%)Total RX: 160767.31Gbits, Total TX: 122118.41Gbits
[1484275464] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (122978293798.13%) too high
[1484275523] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 39.74Mb (1.99%), Average OUT: 66.10Mb (3.30%)Total RX: 160737.74Gbits, Total TX: 122126.34Gbits
[1484276424] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (120962256194.23%) too high
[1484276484] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 24.12Mb (1.21%), Average OUT: 68.91Mb (3.45%)Total RX: 160771.44Gbits, Total TX: 122191.18Gbits
[1484276608] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;1;WARNING - load average: 4.17, 2.89, 2.18
[1484276668] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;OK;SOFT;2;OK - load average: 3.47, 2.95, 2.25
[1484276788] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;1;WARNING - load average: 5.28, 3.70, 2.60
[1484276848] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;2;WARNING - load average: 3.17, 3.43, 2.58
[1484276903] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.61%) too high
[1484276908] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;SOFT;3;WARNING - load average: 3.57, 3.48, 2.65
[1484276964] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 31.26Mb (1.56%), Average OUT: 72.48Mb (3.62%)Total RX: 160786.88Gbits, Total TX: 122223.98Gbits
[1484276968] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;WARNING;HARD;4;WARNING - load average: 4.08, 3.89, 2.86
[1484277148] SERVICE ALERT: OS-MONITOREO;Carga OS-MONITOREO;CRITICAL;HARD;4;CRITICAL - load average: 4.60, 4.08, 3.10
[1484277444] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.79%) too high
[1484277504] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 24.60Mb (1.23%), Average OUT: 632.15Mb (31.61%)Total RX: 160837.87Gbits, Total TX: 122292.42Gbits
[1484277563] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.69%) too high
[1484277624] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 26.51Mb (1.33%), Average OUT: 45.94Mb (2.30%)Total RX: 160807.06Gbits, Total TX: 122298.85Gbits
[1484278043] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.69%) too high
[1484278104] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 26.48Mb (1.32%), Average OUT: 79.72Mb (3.99%)Total RX: 160821.30Gbits, Total TX: 122293.35Gbits
[1484278583] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.06%) too high
[1484278644] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 589.63Mb (29.48%), Average OUT: 64.98Mb (3.25%)Total RX: 160871.98Gbits, Total TX: 122327.19Gbits
[1484278703] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.83%) too high
[1484278764] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 30.57Mb (1.53%), Average OUT: 54.93Mb (2.75%)Total RX: 160841.59Gbits, Total TX: 122334.06Gbits
[1484279123] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671659.93%) too high
[1484279183] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 31.55Mb (1.58%), Average OUT: 62.08Mb (3.10%)Total RX: 160858.21Gbits, Total TX: 122358.91Gbits
[1484279723] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.38%) too high
[1484279783] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 25.21Mb (1.26%), Average OUT: 64.63Mb (3.23%)Total RX: 160875.27Gbits, Total TX: 122396.62Gbits
[1484280263] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.15%) too high
[1484280323] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 31.56Mb (1.58%), Average OUT: 53.31Mb (2.67%)Total RX: 160891.39Gbits, Total TX: 122427.71Gbits
[1484280863] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.81%) too high
[1484280923] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 33.13Mb (1.66%), Average OUT: 57.84Mb (2.89%)Total RX: 160909.59Gbits, Total TX: 122463.40Gbits
[1484281523] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (122978293799.21%) too high
[1484281584] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 28.06Mb (1.40%), Average OUT: 55.98Mb (2.80%)Total RX: 160929.71Gbits, Total TX: 122499.29Gbits
[1484282063] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.27%) too high
[1484282124] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 68.24Mb (3.41%), Average OUT: 61.08Mb (3.05%)Total RX: 160946.35Gbits, Total TX: 122564.89Gbits
[1484282723] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (125062671660.14%) too high
[1484282784] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 20.78Mb (1.04%), Average OUT: 79.41Mb (3.97%)Total RX: 160962.62Gbits, Total TX: 122569.58Gbits
[1484283204] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - OUT bandwidth (120962256195.97%) too high
[1484283263] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 23.76Mb (1.19%), Average OUT: 64.30Mb (3.22%)Total RX: 160973.40Gbits, Total TX: 122599.85Gbits
[1484283383] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];CRITICAL;SOFT;1;CRITICAL - IN bandwidth (125062671658.19%) too high
[1484283444] SERVICE ALERT: FTTH-SABANA01;TRONCAL [AMERICAN_DATA02 - Po5[G1/23 - G1/24]];OK;SOFT;2;OK - Average IN: 20.36Mb (1.02%), Average OUT: 62.47Mb (3.12%)Total RX: 160977.33Gbits, Total TX: 122610.55Gbits